A rapid transfer of virions coated with heparan sulfate from the ECM to CD151 defines an early step in the human papillomavirus infection cascade

Author response:

The following is the authors’ response to the previous reviews

eLife Assessment

This study provides valuable insight into the role of actin protrusions in mediating early pre-endoyctic steps of human papillomavirus entry at the cell surface. Using state-of-the-art microscopy in an immortalized keratinocyte model, the authors present mostly solid evidence that filopodia actively promote the transfer of heparin sulfate-coated virions from the extracullar matrix to the viral entry factor CD151. Remaining gaps in the mechanistic model could be further supported by including a more expansive analysis of the fixed microscopy samples and live cell imaging to distinguish virion transfer from direct binding.

We thank the editorial team for the improved eLife assessment. Regarding the remaining gap, we agree that it is not clear why the large majority of the virions indeed are transferred and not directly binding virions.

Public Reviews:

Reviewer #1 (Public review):

Summary:

The author's goal was to arrest PsV capsids on the extracellular matrix using cytochalasin D. The cohort was then released and interaction with the cell surface, specifically with CD151 was assessed.

The model that fragmented HS associated with released virions mediates the dominant mechanism of infectious entry has only been suggested by research from a single laboratory and has not been verified in the 10+ years since publication. The authors are basing this study on the assumption that this model is correct, and these data are referred to repeatedly as the accepted model despite much evidence to the contrary. The discussion in lines 65-71 concerning virion and HSPG affinity changes is greatly simplified. The structural changes in the capsid induced by HS interaction and the role of this priming for KLK8 and furin cleavage has been well researched. Multiple laboratories have independently documented this. If this study aims to verify the shedding model, additional data needs to be provided.

Comment of the authors: the above paragraph is copied from the very first review and describes the situation before revision.

Note on revisions:

The authors did an excellent job in their revision to include data from the effect of proteolytic priming on their observed virion transfer to the cell body. All other minor issues were addressed adequately.

We are grateful that the referee acknowledges that we addressed all issues adequately.

The work could be especially critical to understanding the process of in vivo infection.

We agree, and would like to point out that a similar comment was raised by the reviewing editor assigned to our original submission, John Schiller. For unknown reasons, he was no longer involved in the evaluation of the revision.

Reviewer #2 (Public review):

The study design involves infecting HaCaT cells (immortalised keratinocytes mimicking basal cells of a target tissue) and observing virus localization with and without actin polymerization inhibition by cytochalasin D (cytoD) to analyze virion transfer from the ECM to the cell via filopodial structures, using cellular proteins as markers.

In the context of the model system, the authors stress in the revised version the importance of using HaCaT cells as a relevant 'polarized' cell model for infection. The term 'polarized' is used in the cell biological literature for epithelial cells to describe a strict apical vs. basolateral demarcation of the plasma membrane with an established diffusion barrier of the tight junction. However, HaCat cells do not form tight junctions. In squamous epithelia, such barriers are only found in granular layers of the epithelium. The published work cited in support of their claims either does not refer to polarity or only in the context of other cells such as CaCo-2 cells.

We thank the reviewer for this important clarification and fully agree. HaCaT cells do not form tight junctions and therefore do not fulfill the classical definition of polarized epithelial cells with a strict apical basolateral diffusion barrier. In response to this comment, we have removed the term “polarized” in reference to HaCaT cells throughout the revised manuscript. Our intention was not to imply classical epithelial polarity, but rather to emphasize that HaCaT cells represent a functionally relevant keratinocyte model that recapitulates key early steps of HPV infection observed in vivo, particularly abundant ECM deposition enabling for strong virion binding to the ECM.

We now state on line 120: “PsVs that bind to the ECM at sites distal from the cell body are unable to establish direct contact with entry receptors, until the cell migrates onto them or they are transported along cell protrusions towards the cell body (Schelhaas et al., 2008; Smith et al., 2008). Both cell migration and protrusion transport depend on actin dynamics (Schaks et al., 2019). We aimed for blocking these active recruitment mechanisms in HaCaT cells, a cell line that is widely used as a cell culture model for HPV infection. They resemble primary keratinocytes in several key aspects: they are not virally transformed and produce large amounts of ECM, promoting interactions between viruses and ECM components and thereby facilitating infection (Bienkowska-Haba et al., 2018; Gilson et al., 2020). In addition, subconfluent HaCaT cells form filopodia and filopodial transport is used for the recruitment of ECM-bound virus particles to the cell body (Schelhaas et al., 2008, Smith et al., 2008). Together, these features make HaCaT cells a suitable model for studying active PsV recruitment from the ECM to the cell surface.”

Overall, the matter of polarity would be important, if indeed the virus could only access cell-associated HSPGs as primary binding receptor, or the elusive secondary receptor via the ECM in the used model system (HaCaT cells), if they would locate exclusively basolaterally.

We apologize for not having stressed enough that virions bind as well directly to the not imaged, upper cell membrane. To make clear that HaCaT cells are still a suitable model for studying active recruitment, throughout the manuscript, we worked on the following issues (this is an outline, for details see below):

(1) We now discuss adequately that virions reach cell surface receptors either by passive diffusion or by active transport mechanisms, the latter involving actin dynamics (filopodial transport and cell migration), to which we refer in the revised manuscript as active recruitment.

(2) We explain why the large majority of virions in the microscopic assay are actively recruited virions.

(3) We explain the difference between biochemical infection assays that do not differentiate between passive and active recruitment, and microscopic assays studying the basal cell membrane and by this primarily actively recruited virions

This is at least not the case for binding, as observed in several previous publications (just two examples: Becker et al, 2018, Smith et al., 2008). With only a rather weak attempt at experimental verification of their model system with regards to polarity of binding, the authors then go on to base their conclusions on this unverified assumption.

We agree with the reviewer that strict epithelial polarity would only be relevant if HPV binding or receptor accessibility were confined to the basolateral membrane, which is not the case in HaCaT cells, as shown previously (e.g., Becker et al., 2018; Smith et al., 2008). However, our conclusions do not rely on strictly polarity-dependent binding.

We added the following paragraphs clarifying that (i) in HaCaT cells PsVs also bind by passive diffusion to the upper cell membrane and that (ii) at the basal membrane the large majority of imaged PsVs is actively recruited.

Line 332: “…, the lower PCC at 0 min/CytD suggests that without active recruitment less PsVs reach CD151. At 30 min after CytD, the PCC has reached the level of 0.1 as in the control, which is in line with the idea of fast recruitment as observed in Figure 4. To follow how the basal cell membrane is populated with PsVs over time, as additional analysis we determined the PsVs per µm2 in ROIs placed in the cell body region. At 0 min, CytD reduces the PsV density to 19 - 33%, albeit the effect is not significant, and at 180 min/CytD the same PsV density as in the control is reached (Supplementary Figure 6A and B). Overall, under CytD there was a trend towards less PsVs present (Supplementary Figure 6A and B). Hence, both Figure 5C and Supplementary Figure 6A and B suggest that active virion transport is required to reach efficiently the basal membrane.”

Line 447: “Throughout all experiments, we observe at 0 min/CytD only few PsVs at the basal membrane (Figure 1A, Supplementary Figure 6A and B; see also PCC at 0 min between PsVs an CD151 in Figure 5C), suggesting that in the absence of active recruitment the access to the basal membrane via passive diffusion is limited. We wondered, how many PsVs may bind to the cell membrane without a diffusion barrier? For this reason, we incubated EDTA detached HaCaT cells in suspension with PsVs for 1 h at 4 °C, followed by re-attachment for 1 h. Under these conditions, we find, despite of a shorter incubation time (1 h versus 5 h), a roughly 3-fold larger PsV density (1.7 PsVs/µm2 (Supplementary Figure 6D)) than the highest density observed in the other experiments. However, it should be noted that values of the different experiments cannot be directly compared. Aside from the different treatments, another difference lies in the size of the imaged membrane. The re-attachment of cells is not complete after 1 h (compare size of adhered membranes in Supplementary Figure 6A and 1A), wherefore the membranes are likely strongly ruffled, which results in the underestimation of the membrane area. As a result, we overestimate the PsVs per µm2 adhered membrane (please note that we cannot re-attach cells for longer times as we then lose PsVs due to endocytosis). In any case, the experiment suggests that PsVs bind more efficiently to membrane surface receptors without a diffusion barrier. We conclude that in our assay PsVs cannot readily bypass the active PsV recruitment by diffusing directly to the basal cell membrane, which is plausible, because to make this happen a 55 nm large PsV must diffuse through the narrow gap between glass-coverslip and adhered cell.”

Line 538: “The analyzed PsVs hardly bind to the basal cell surface directly by diffusion (Supplementary Figure 6, compare PsV maxima density at 0 min/CytD in A and B to C). Therefore, the actin-driven virion transport would play a decisive role in HPV infection if cells would form a monolayer with a disruption at which ECM is present and that is approached by PsVs, a scenario similar to in vivo infection. In addition, cell migration could establish contact between PsVs and the cell surface.”

Line 548: “…that can readily bind to the upper cell membrane. We are not aware of a PsV translocation mechanism from the upper to the basal membrane. Therefore, in our assay, PsVs bound to the upper membrane are not expected to show up at the basal membrane. Comparing 0 min of control and CytD (Supplementary Figure 6A and B), we find that compared to the control 19 - 33% of the PsVs reach the basal membrane in the absence of active transport, or in other words, most likely by passive diffusion. Actually, the range from 19 – 33% must be a strong overestimate as PsVs in the control are in transit and many actively recruited PsVs are already internalized during the 5 h incubation period. For this reason, we propose that most likely much less than 10% of the PsVs reach the basal membrane by diffusion. Moreover, in the absence of the diffusion barrier, the density of bound PsVs is strongly increased (Supplementary Figure 6D), showing indirectly that at the basal membrane the binding sites are difficult to access without active recruitment. Taken together, we propose the large majority of PsVs analyzed in our assay are ECM bound and actively recruited to the basal cell membrane.”

This is one example of several in the manuscript, where claims for foundational premises, observations, and/or conclusions remain undocumented or not supported by experimental data.

Another such example is the assumption of transfer of the virus from ECM to the tetraspanin CD151. Here, the conclusions are based on the poorly documented inability of the virus to bind to the cell body, which is in stark contrast to several previous publications, and raises questions.

We hope with the above changes we made clear that virions can also directly bind to the cell body. We also added a paragraph discussing differences between biochemical and microscopic assays.

Line 568: “In this scenario, sub-confluent HaCaT cells, or even better single HaCaT cells, would be an ideal model system for the microscopic study of these very early infection steps that involve ECM attachment and subsequent active recruitment, as supposed to occur during in vivo infection of basal keratinocytes after binding of virions to the basement membrane (Bienkowska-Haba et al., 2018; Day and Schelhaas, 2014; Kines et al., 2009; Schiller et al., 2010). In contrast, in biochemical infection assays, virions diffusing to HSPGs on the cell surface, and by this bypassing active recruitment, are assayed together with the actively recruited virions. Should cells secrete little ECM and are grown to confluency, the passively binding virions are supposed to strongly dominate the infection rate in a biochemical infection assay.”

There are a number of important additional issues with the manuscript:

First, none of the inhibitors have been tested in their system for efficacy and specificity, but rely on published work in other cell types. This considerably weakens the confidence on the conclusion drawn by the authors.

We use inhibitors CytD, blebbistatin, leupeptin and furin inhibitor I. The below references are examples reporting the usage of the inhibitors on HaCaT cells studied in the context of HPV infection.

Furin inhibitor I:

Cruz et al., Cleavage of the HPV16 Minor Capsid Protein L2 during Virion Morphogenesis Ablates the Requirement for Cellular Furin during De Novo Infection. Viruses, 2015; doi.org/10.3390/v7112910

Cytochalasin D/Blebbistatin:

Schelhaas et al., Human papillomavirus type 16 entry: retrograde cell surface transport along actinrich protrusions. PLoS Pathog., 2008. doi: 10.1371/journal.ppat.1000148.

Smith et al., Virus activated filopodia promote human papillomavirus type 31 uptake from the extracellular matrix. Virology, 2009; doi.org/10.1016/j.virol.2008.08.040 and

Leupeptin/Furin inhibitor I:

Cerqueira et al., Kallikrein-8 Proteolytically Processes Human Papillomaviruses in the Extracellular Space To Facilitate Entry into Host Cells. J. Virology, 2015; doi.org/10.1128/jvi.00234-15

Moreover, the reversible inhibitory effect of CytD the key inhibitor, used in this study on transport and infection is validated in this study. However, we discuss this data now in the context of directly binding virions more critically.

Line 485: “Hence, the infection assay suggests that the treatment is largely reversible and only slightly harmful, if at all. However, the luciferase infection assay does not distinguish between actively recruited PsVs and PsVs that bind passively by diffusion to the upper membrane. The latter fraction likely dominates the total infection rate and should be less affected by CytD than the fraction of actively recruited PsVs. Therefore, if the infection pathway of a small fraction of actively recruited PsVs is irreversibly inhibited, we may not be able to detect this effect on the background of unaffected passively binding PsV.”

Second, the authors aim to study transfer from ECM to the cell body and effects thereof. However, there are still substantial amounts of viruses that bind to the cell body compared to ECM-bound viruses in close vicinity to the cells.

Regarding direct binding to the cell body, please see our detailed reply above.

This is in part obscured by the small subcellular regions of interest that are imaged by STED microscopy, or by the use of plasma membrane sheets. This remains an issue despite the added Supple. Fig. 1, where also only sub cellular regions are being displayed. As a consequence the obtained data from time point experiments is skewed, and remains for the most part unconvincing, largely because the origin of virions in time and space cannot be taken into account. This is particularly important when interpreting the association with HS, the tetraspanin CD151, and integral alpha 6, as the low degree of association could be originating from cell bound and ECM-transferred virions alike.

We hope with the above explanations it is plausible that the imaged virions primarily reach the basal membrane by active recruitment.

Third, the use of fixed images in a time course series also does not allow to understand the issue of a potential contribution of cell membrane retraction upon cytoD treatment due to destabilisation of cortical actin. Or, of cell spreading upon cytoD washout. The microscopic analysis uses an extension of a plasma membrane stain as marker for ECM bound virions, this may introduce a bias and skew the analysis.

The referee is correct in pointing out that cell spreading after CytD wash off would affect our analysis, e.g. by increasing the overlap between PsVs and the cell body although no active recruitment via filopodial transport and cell migration occurs. An argument speaking against this possibility is the lack of increase in the PCC between PsVs and F-actin after CytD removal, if the protease inhibitor leupeptin was present (Figure 2B and D). Leupeptin prevents PsV/phalloidin overlap despite restored actin polymerization after washout of both inhibitors, suggesting that priming is required for increased PsV–actin association and is too slow to change PCC within 60 min. These results support that the observed overlap reflects active, priming-dependent recruitment rather than cell morphology changes.

We state on line 252: “Moreover, the experiment suggests that without PsV priming the PCC between PsV-L1 and F-actin does not increase, for instance, due to cell spreading after CytD removal.”

On line 494, we state “However, we assume that this is rather unlikely, as cell spreading would increase the PCC between PsVs and F-actin under a condition where PsVs are not-primed (and therefore not actively recruited) but cell spreading occurs, which is not the case in Figure 2B and D (CytD/leupeptin).”

Fourth, while the use of randomisation during image analysis is highly recommended to establish significance (flipping), it should be done using only ROIs that have a similar density of objects for which correlations are being established. For instance, if one flips an image with half of the image showing the cell body, and half of the image ECM, it is clear that association with cell membrane structures will only be significant in the original. But given the high density of objects on the plasma membrane, I am not convinced that doing the same by flipping only the plasma membrane will not also obtain similar numbers than the original.

Regarding the association of PsVs with CD151 and HS, we corrected for random background with reference to a calibration line that describes the random background association in dependence of the density of objects. We now refer to this issue on line 343: “…, the fraction of PsVs closely associated with CD151 is around 10% (Figure 5D, control), after correction for random background association, for which we used a calibration line based on the same density of PsVs in flipped images (see Supplementary Figure 7).”

In the legend of Supplementary Figure 7 we state: “…The fraction of closely associated PsVs (PsV-L1 maxima with a distance ≤ 80 nm to the next nearest CD151 maximum) in the Control of Figure 5 was analyzed on original and flipped images (for an example of a flipped image see Supplementary Figure 5A)…on flipped images, we often find values more than half of the values of the original images, demonstrating that many PsVs have a distance ≤ 80 nm to CD151 merely by chance, in the following referred to as background association…We take the altogether 24 fraction values obtained on flipped images (12 values from Control and CytD each), and plot the fraction of closely associated PsVs against the average CD151 maxima density in the respective images. As can be seen in (C), the fraction increases with the maxima density, as the chance of a distance ≤ 80 nm increases with the maxima density. The fitted linear regression line describes how the background association depends from the maxima density. As a result, the background association (y) can be calculated for any maxima density (x) with the equation y = 2.04 • x. The CytD/0 min condition may be overcorrected, if it includes many images with CD151 flipped onto peripheral PsVs that actually are distal to CD151 (for an example ROI see Supplementary Figure 5A). On the other hand, PsVs right at the cell border, where CD151 staining tends to be strong (Supplementary Figure 5A), after flipping have less CD151 than before, contributing to undercorrection.”

When omitting the CytD/0 min values, we obtain essentially the same calibration line.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

There are further issues that are not pertaining to the study design that I find important.

Fig.1

There are few, if any, filopodia in untreated cells. It would be good to quantify their abundance to substantiate that resting HaCat cells are indeed a good model for filopodial transport bs. membrane retraction / spreading.

We see filopodia in untreated HaCaT cells (although quite variable in abundance, please see control cells in e.g. Figure 3 and 8 and Supplementary Figure 2).

In HaCat ECM the virus binds also to laminin-332 for a good part. Would this not also confound the analysis?

We agree with the reviewer that in HaCaT-derived ECM, virus binding is not restricted to heparan sulfate (HS), and that laminin-332 represents an additional relevant binding partner. Indeed, viruses bound to laminin-332 may likewise be transported toward the cell body via laminin-binding integrins. We therefore consider laminin-332 to act as a parallel attachment factor alongside HS rather than as a mutually exclusive alternative.

However, the primary aim of this study was not to comprehensively map all ECM binding partners, but to analyze the actin-dependent transport of ECM-bound virus particles. HS was chosen as a representative and well-characterized ECM marker for initial virus attachment. Importantly, inhibition of actin dynamics by cytochalasin D blocks this transport process downstream of initial binding. Thus, irrespective of whether the virus is initially bound to HS, laminin-332, or both, the readout reflects interference with the same actin-dependent transport mechanism.

Consequently, the presence of laminin-332 binding does not confound our analysis, as the experimental outcome is determined by inhibition of transport rather than by the specific ECM attachment factor. Nonetheless, we acknowledge laminin-332 as an important parallel interaction partner and had already mentioned it the first version of the manuscript, but removed the sentence during the last revision, that has now been added again. On line 593 we state: “Finally, not all PsVs bound to the ECM are expected to bind to HS but could also bind to laminin 332 (Culp et al., 2006).”

Fig.2

Would benefit from live cell analysis. There are considerable amounts of virions on the cell body, which partially contradicts statements from Fig. 1. The fast transfer to the cell body after cyto D washout is based on the assumption that filopodia formation and transport along them (and not membrane extension) occurs quickly. Is this reasonable? Does membrane extension and migration occur between 0 min and later time points?

Regarding membrane extension after CytD removal, that in the analysis may be indistinguishable from active recruitment transfer, please see our reply above (no PCC increase between PsV-L1 and F-actin after CytD removal if leupeptin is employed). Regarding migration, we now included this possibility as an active recruitment mechanism that may occur in parallel to filopodial transport (please see our reply above).

Fig.4

How are the subcellular ROIs chosen? Is there not a bias by not studying a full cell?

In Figure 4 we are specifically interested in the time course of PsV diminishment from the cell periphery. The ROIs are generated with reference to the membrane staining, using the cell body delineation as a starting point. For details about how ROIs are generated, please see legend of Figure 4 and materials and methods.

Fig. 5/6

The data needs a better analysis on correlation by using randomisation as explained above.

Please see our reply above. The association between PsVs and CD151 or HS has been corrected using a calibration line based on the same density of objects.

Fig. 8. Why does blebbistatin block the transport only partially? Previous work on actin retrograde flow suggests that in the absence of myosin II function the transport stops completely. Would this not be a concern, when interpreting the city D data?

Is the referee referring to Schelhaas et al., 2008 that we cite in the paper? In this paper, in HeLa cells blebbistatin reduced the directed particle motion by 82%, but not completely.

Suppl. Fig. 1A, B: Intented to adress the issue of viruses binding to the cell body, it unfortunately falls short. It would have been better to analyse complete cells rather than ROIs, or better even, a comprehensive analysis of cell islets (boundary cells vs. central cells, with cell body to cell periphery).

This experiment addresses the increase in PsV density resulting from active recruitment. Outlining entire cells would include also PsVs close to the cell edge that have not been actively recruited.

Regarding cell islets (we call them patches of confluent cells as islets may be confused with e.g. more structured Langerhans islets), there are hardly any PsVs at the basal membrane. We state on line 135: “Frequently, we observe patches of confluent cells which are common to HaCaT cells. Cells at the center of these patches are dismissed during imaging, because hardly any PsVs are bound to their basal membrane, indicating that PsVs do rather not reach this area by passive diffusion. Instead, we focus on isolated HaCaT cells or cells at the periphery of cell patches. At these cells, we find more PsVs per cell than one would expect from the employed ≈ 50 viral genome equivalents (vge) per cell, indicating that PsVs are unequally distributed between the cells.”

Is the difference between untreated and cytoD treated significant?

We stated in the Figure legend that the difference is not significant (the exact p value is p = 0.089). We now have revised the Figure (previously Supplementary Figure 1A and B, now Supplementary Figure 6A and B), showing the PsV density at the basal membrane over time, also for the experiment shown in Figure 6. The now revised Figure (Supplementary Figure 6A and B) is discussed together with the re-attachment experiment (Supplementary Figure 6C and D), in order to compare the PsV accessibility to the cell membrane with and without diffusion barrier. Please see our reply above (paragraph starting at line 447).

Changes from the earlier version to the revised version

Abstract

Earlier version: Human Papillomaviruses (HPVs) are the underlying cause of several types of cancer, albeit they are mostly known for their association with cervical carcinoma. The virions must reach their target cells through a break in the epithelial barrier. After binding to heparan sulfate (HS) of the extracellular matrix (ECM), they translocate to the cell surface and co-internalize with the entry factor CD151.

Revised version: Human Papillomaviruses (HPVs) are the underlying cause of several types of cancer, albeit they are mostly known for their association with cervical carcinoma. The virions reach their target cells through a break in the epithelial barrier. After binding to heparan sulfate (HS) of the extracellular matrix (ECM), they are recruited via actin-dependent mechanisms to the cell surface where they co-internalize with the entry factor CD151.

Earlier version: The in vivo occurring translocation from the ECM to entry factors at the cell membrane may be bypassed in nonpolarized cells with low ECM secretion. To specifically investigate these early events of the infection cascade, we use polarized keratinocytes. They produce a robust ECM and have a restricted expression of entry receptors primarily to the substrate-adhered basal membrane. Virion access to the basal membrane by diffusion is strongly limited, such that translocation from the ECM is the main pathway used for infection. We block the translocation from ECM attachment sites to the cell body, release the blocking, and monitor the association of virions with CD151 or HS. We observe quick virion translocation from the ECM to the cell body within 15 min. During translocation, virions associate with the tetraspanin CD151 present at the cell border or at filopodia. Translocating virions are decorated with HS, which they lose in the next few hours, presumably prior to endocytosis.

Revised version: The in vivo occurring active recruitment from the ECM to the cell surface may be bypassed in cell culture where virions reach the cell surface simply by passive diffusion. To specifically investigate these early events of the infection cascade, we use HaCaT keratinocytes as they produce a robust ECM enabling for abundant virion binding to ECM components such as HS before transfer to cell surface receptors and infection. Employing microscopy, we focus on the basal membrane that for virions is difficult to access by diffusion. We block the active recruitment from ECM attachment sites to the cell body, release the blocking, and monitor the association of virions with CD151 or HS. We observe quick virion recruitment from the ECM to the cell body within 15 min. During recruitment, virions associate with the tetraspanin CD151 present at the cell border or at filopodia. These virions are decorated with HS, which they lose in the next few hours, presumably prior to endocytosis.

Our observations reveal a rapid step in the HPV infection cascade: the transfer of HS-coated virions from the ECM to CD151. This step is too fast to account for the asynchronous uptake of HPVs which is likely driven by glycan- and capsid processing.

Introduction

Already in the 80s, Harald zur Hausen proposed a role of human papillomaviruses in cancer (zur Hausen, 2009). Since then, five more classes of oncogenic viruses have been identified (Galati et al., 2024a). To date, it is assumed that more than 10% of the worldwide human cancer burden is associated with infectious agents (Galati et al., 2024a), from which about a half is caused by Papillomaviridae (Martel et al., 2017). Thus, the understanding of viral entry strategies has implications going beyond the classical treatment of acute viral infections.

Earlier version: Human papillomaviruses are small, non-enveloped viruses with a diameter of ≈55 nm. The icosahedral capsid is mainly composed of pentameric L1 capsomers. Together with less L2 capsid proteins, capsomers surround a histone core bearing a circular double-stranded DNA (Baker et al., 1991; Ozbun and Campos, 2021). From the more than 200 phylogenetically classified HPV genotypes, the most oncogenic ones are HPV16 and HPV18 (Galati et al., 2024b), which are responsible for about 70 - 80% of the cervical cancer cases (Christiansen et al., 2015). In addition, they cause other severe cancers such as anogenital, head and neck tumors (Doorbar et al., 2012).

Revised version: Human papillomaviruses are small, non-enveloped viruses with a diameter of ≈55 nm. The icosahedral capsid is mainly composed of pentameric L1 capsomers. Together with fewer L2 capsid proteins, capsomers surround a histone core bearing a circular double-stranded DNA (Baker et al., 1991; Ozbun and Campos, 2021). From the more than 200 phylogenetically classified HPV genotypes, the most oncogenic ones are HPV16 and HPV18 (Galati et al., 2024b), which are responsible for about 70 - 80% of the cervical cancer cases (Christiansen et al., 2015). In addition, they cause other severe cancers such as anogenital, head and neck tumors (Doorbar et al., 2012).

Earlier version: For papillomavirus infection, a disruption of the epithelial barrier is a prerequisite, through which virions reach mitotically active basal cells of the epithelia (Ozbun, 2019). Here, virions bind to the linear polysaccharide heparan sulfate (HS) that is present in the extracellular matrix (ECM) but as well on the plasma membrane surface. HS is attached to proteins forming so called heparan sulfate proteoglycans (HSPGs). Positively charged and polar residues of the L1 capsid protein form multiple heparan sulfate binding sites that interact with negative charges of HS, resulting in a strong bond and virus uptake into the infectious pathway (Dasgupta et al., 2011; Giroglou et al., 2001; Joyce et al., 1999; Knappe et al., 2007; Surviladze et al., 2015). While in cell culture virions bind to HS on both the cell surface and the ECM, it has been suggested that in vivo they bind predominantly to HS of the extracellular basement membrane (Day and Schelhaas, 2014; Kines et al., 2009; Ozbun and Campos, 2021; Schiller et al., 2010). In any case, the tight bond between the linear polysaccharide and virions must be disrupted before they can bind to a yet unknown secondary receptor on the cell surface, followed by internalization (Ozbun and Campos, 2021).

Revised version: For papillomavirus infection, a disruption of the epithelial barrier is a prerequisite, through which virions reach mitotically active basal cells of the epithelia (Ozbun, 2019). Here, virions bind to the linear polysaccharide heparan sulfate (HS) that is present in the extracellular matrix (ECM) and at the plasma membrane surface. HS is covalently linked to proteins forming so called heparan sulfate proteoglycans (HSPGs). Positively charged and polar amino acid residues of the L1 capsid protein form multiple heparan sulfate binding sites that interact with negative charges of HS, resulting in a strong bond (Dasgupta et al., 2011; Giroglou et al., 2001; Joyce et al., 1999; Knappe et al., 2007; Surviladze et al., 2015). While in cell culture virions bind to HS on both the cell surface and the ECM, it has been suggested that in vivo they bind predominantly to HS of the extracellular basement membrane (Day and Schelhaas, 2014; Kines et al., 2009; Ozbun and Campos, 2021; Schiller et al., 2010). In any case, the link between the linear polysaccharide and virions must be disrupted before they can bind to a yet unknown secondary receptor on the cell surface, followed by internalization (Ozbun and Campos, 2021).

Earlier version: The strong electrostatic bonding precludes HS release via dissociation. Two release mechanisms are discussed, that mutually are not exclusive. In the so-called priming model, binding to HS structurally results in capsid enlargement and softening (Feng et al., 2024), followed by the exposure and cleavage of L1 by kallikrein-8 (KLK8) (Cerqueira et al., 2015), exposure of L2 by cyclophilin (Bienkowska-Haba et al., 2009), and L2 cleavage by furin (Richards et al., 2006). After these structural modifications of the capsid surface the ‘primed’ virion is able to bind to the secondary receptor. In an alternative model, the HS-virion bond remains. However, cleavage of HS/HSPGs by heparanases and proteinases produces HS fragments. Therefore, albeit still bound to the virion surface, the fragmented HS no longer anchors the virion to the ECM (Surviladze et al., 2012; Surviladze et al., 2015).

Revised version: The strong electrostatic bonding precludes dissociation as a virion release mechanism. Two alternatives are discussed, that mutually are not exclusive. In the so-called priming model, binding of HS to the capsid results in capsid enlargement and softening (Feng et al., 2024), followed by the exposure and cleavage of L1 by kallikrein-8 (KLK8) (Cerqueira et al., 2015), exposure of L2 by cyclophilin (Bienkowska-Haba et al., 2009), and L2 cleavage by furin (Richards et al., 2006). After these structural modifications of the capsid surface the so-called ‘primed’ virion is able to bind to the secondary receptor. In an alternative model, the HS-virion bond persists. However, heparanases and proteinases cleave HS/HSPGs into fragments. As a result, albeit still bound to the virion surface, the now fragmented HS/HSPG no longer anchors the virion to the ECM (Surviladze et al., 2012; Surviladze et al., 2015). Next, released virions may reach cell surface receptors simply by passive diffusion. However, active recruitment is possible as well. For instance, PsVs migrate along actin-rich protrusions from the ECM towards the cell body (Schelhaas et al., 2008; Smith et al., 2008). Moreover, cells may reach virions by migrating onto an ECM lawn decorated with virions.

Earlier version: Once the virions are released, they could reach the cell surface simply by passive diffusion. In a co-culture trans-well assay, pseudovirions (PsVs) from donor cells infect spatially separated receiver cells (Surviladze et al., 2012), demonstrating that free diffusion of virions may be sufficient for infection. However, an active transport mechanism has been reported as well. PsVs migrate along actin-rich protrusions from the ECM towards the cell body (Schelhaas et al., 2008; Smith et al., 2008). While it is generally agreed on that HS is a primary virion attachment site, the molecular identity of the secondary receptor on the main cell body surface is unknown. The secondary complex is most likely a multimeric complex rather than a single molecular component. Possible candidates are proteins crucial for cell entry, as the tetraspanin CD151 (Mikuličić et al., 2019; Scheffer et al., 2013; Spoden et al., 2008), integrin-α6 (Itgα6) (Evander et al., 1997; Yoon et al., 2001), growth factor receptors (Mikuličić et al., 2019; Surviladze et al., 2012), and the annexin A2 heterotetramer (Dziduszko and Ozbun, 2013; Woodham et al., 2012). From these molecules, the tetraspanin CD151 could play a coordinating role, as tetraspanins organize viral entry platforms in many types of viral infections, including infections with coronavirus, cytomegalovirus, hepatitis C virus, human immunodeficiency virus, human papilloma virus, and influenza virus (Bruening et al., 2018; Earnest et al., 2015; Florin and Lang, 2018; Hantak et al., 2019; Hochdorfer et al., 2016; Scheffer et al., 2013; Zona et al., 2013). Tetraspanin entry platforms could form in a slow and stochastic fashion, which would provide an explanation for the asynchronous virion uptake with half-times of above 10 h (Becker et al., 2018). However, it is unclear whether virions associate with CD151 already in the moment of virion transfer from the ECM to the cell surface or afterwards in preparation of the endocytic uptake.

Revised version: While it is generally agreed on that HS is a primary virion attachment site, the molecular identity of the secondary receptor on the main cell body surface is unknown. The secondary receptor complex is most likely of multimeric nature rather than a single molecular component. Possible candidates are proteins known to be crucial for cell entry, as the tetraspanin CD151 (Mikuličić et al., 2019; Scheffer et al., 2013; Spoden et al., 2008), integrin-α6 (Itgα6) (Evander et al., 1997; Yoon et al., 2001), growth factor receptors (Mikuličić et al., 2019; Surviladze et al., 2012), and the annexin A2 heterotetramer (Dziduszko and Ozbun, 2013; Woodham et al., 2012). From these molecules, the tetraspanin CD151 could play a coordinating role, as tetraspanins organize viral entry platforms in many types of viral infections, including infections with coronavirus, cytomegalovirus, hepatitis C virus, human immunodeficiency virus, human papilloma virus, and influenza virus (Bruening et al., 2018; Earnest et al., 2015; Florin and Lang, 2018; Hantak et al., 2019; Hochdorfer et al., 2016; Scheffer et al., 2013; Zona et al., 2013). Such tetraspanin entry platforms could form in a slow and stochastic fashion, which would provide an explanation for the asynchronous virion uptake with half-times of above 10 h (Becker et al., 2018). However, it is unclear whether virions associate with CD151 already in the moment of virion transfer from the ECM to the cell surface, or in a subsequent step, e.g. in preparation of endocytosis.

Earlier version: In this study, we explore the relationship between actin-dependent transport and CD151 association in HPV infection. We employ the cell permeable mycotoxin and actin polymerization inhibitor cytochalasin D (CytD). We find that CytD causes the trapping of HPV16 PsVs in the ECM of a polarized human keratinocyte cell line (HaCaT cells), where they remain accumulated adjacent to the cell periphery. This blocking of the virus transfer is accompanied by co-accumulation of HS in the ECM area. Upon CytD removal, HS-decorated PsVs move from the ECM to the cell body and CD151. The association of PsVs with CD151 persists within the next few hours, whereas the HS coat is stripped off. These findings distinguish a step in the infection cascade with CD151 playing an early role directly after the release of virions from the ECM.

Revised version: Here, we explore the question of a possible link between active virion recruitment to the cell surface and CD151 association. We employ the cell permeable mycotoxin and actin polymerization inhibitor cytochalasin D (CytD). Using a human keratinocyte cell line (HaCaT cells), we find that CytD preserves HPV16 PsVs in the ECM, noticed as PsV accumulations adjacent to the cell periphery. This blocking of the virus transfer is accompanied by co-accumulation of HS in the ECM area. Upon CytD removal, HS-decorated PsVs get from the ECM to the cell body where they associate with CD151. The association of PsVs with CD151 persists within the next few hours, whereas the HS coat is stripped off, and CD151 is observed to agglomerate. These findings distinguish an early step in the infection cascade, the association with CD151 in the moment the virion establishes contact to the cell surface.

Results

Earlier version: Cytochalasin D arrests the translocation of HPV16 PsVs from the extracellular matrix to the cell body surface

Revised version: Cytochalasin D arrests the active recruitment of HPV16 PsVs from the extracellular matrix to the basal cell membrane

Earlier version: The molecular surface of PsVs is immunologically indistinguishable from HPV virus particles (Ozbun and Campos, 2021), which makes them a widely used tool for studying host cell entry. In this study, we employ HPV16 pseudovirions with an encapsidated luciferase reporter plasmid under the control of the HPV16 promotor (Wüstenhagen et al., 2018). Hence, instead of viral DNA a luciferase encoding plasmid enters the nucleus, enabling the analysis of the infection rate via the luciferase activity. Moreover, the plasmids are composed of nucleotides to which fluorophores can be coupled by click-chemistry. This allows for their detection by fluorescence microscopy throughout the entire infection pathway.

Revised version: The molecular surface of PsVs is immunologically indistinguishable from HPV virus particles (Ozbun and Campos, 2021), which makes them a widely used tool for studying host cell entry. In this study, we employ HPV16 pseudovirions with an encapsidated luciferase reporter plasmid under the control of the HPV16 promotor (Wüstenhagen et al., 2018). Hence, instead of viral DNA a luciferase encoding plasmid enters the cell, enabling the analysis of the infection rate via the luciferase activity. Moreover, the plasmids are composed of nucleotides to which fluorophores can be coupled by click-chemistry. This allows for their microscopic detection by fluorescence microscopy, alternatively to immunolabeling of the L1 capsid protein.

Earlier version: Many PsVs that bind to the ECM may locate distal from the cell surface and are thus unable to establish direct contact with entry receptors. However, they are capable of migrating by an actin-dependent transport along cell protrusions towards the cell body (Schelhaas et al., 2008; Smith et al., 2008). We aimed for blocking this transport in HaCaT cells, a cell line that is widely used as a cell culture model for HPV infection. HaCaT cells closely resemble primary keratinocytes in key aspects: they are not virally transformed and produce large amounts of ECM that facilitates infection (Bienkowska-Haba et al., 2018; Gilson et al., 2020). In addition, HaCaT cells exhibit cellular polarity that enforces binding of virus particles to the ECM, as the virions cannot bind to receptors/entry components, such as CD151, Itgα6 and HSPGs that co-distribute on the basolateral membrane of polarized keratinocytes (Cowin et al., 2006; Mertens et al., 1996; Sterk et al., 2000), making them inaccessible by diffusion.

Revised version: PsVs that bind to the ECM at sites distal from the cell body are unable to establish direct contact with entry receptors, until the cell migrates onto them or they are transported along cell protrusions towards the cell body (Schelhaas et al., 2008; Smith et al., 2008). Both cell migration and protrusion transport depend on actin dynamics (Schaks et al., 2019). We aimed for blocking these active recruitment mechanisms in HaCaT cells, a cell line that is widely used as a cell culture model for HPV infection. They resemble primary keratinocytes in several key aspects: they are not virally transformed and produce large amounts of ECM, promoting interactions between viruses and ECM components and thereby facilitating infection (Bienkowska-Haba et al., 2018; Gilson et al., 2020). In addition, subconfluent HaCaT cells form filopodia and filopodial transport is used for the recruitment of ECM-bound virus particles to the cell body (Schelhaas et al., 2008, Smith et al., 2008). Together, these features make HaCaT cells a suitable model for studying active PsV recruitment from the ECM to the cell surface.

Earlier version: We incubate cells for 5 h with PsVs in the absence or presence of 10 µg/ml CytD. Previously, it has been shown that this CytD concentration stops cell migration after a few minutes (Peng et al., 2011). Frequently, we observe patches of confluent cells which are common to HaCaT cells. Cells at the center of these patches are dismissed during imaging, because there are no anterogradely migrating PsVs at these cells. A second reason for our dismissal of these cells is that hardly any PsVs are bound to them, possibly because their basal membranes are inaccessible by diffusion. Instead, we focus on isolated HaCaT cells or cells at the periphery of cell patches. At these cells, we find more PsVs per cell than one would expect from the employed 50 viral genome equivalents (vge) per cell, indicating that PsVs are unequally distributed between the cells. Moreover, these PsVs usually are not homogenously distributed around the cell but concentrate at one region. We investigate the translocation of PsVs from these regions, defining ROIs for analysis that cover PsVs at the periphery and the cell body (see Supplementary Figures 6A and 8A). After the 5 h incubation, cells were fixed and stained by a membrane marker, the dye TMA-DPH, to mark the main cell body (Figure 1A, gray lookup table (LUT)). Additionally, they were stained for PsVs (by click-chemistry; magenta LUT), HS (antibody staining; cyan LUT), and the membrane protein Itgα6 (antibody staining; green LUT). In the control, PsVs locate mostly in the cell body area (Figure 1A, upper row; see white outline that is based on TMA-DPH staining). In contrast, when CytD was present, the cell body area is largely devoid of PsVs (Figure 1A, lower row; please note that the large central bright area is mainly caused by autofluorescence). Instead, PsVs often accumulate adjacent to the cell body. This adjacent region that is up to several µm wide and rich in HS, likely contains the ECM. We conclude that, upon inhibition of actin-dynamics by CytD, PsVs remain accumulated in the ECM area.

Revised version: We incubate cells for 5 h with PsVs in the absence or presence of 10 µg/ml (19.7 µM) CytD, a concentration that stops cell migration after a few minutes (Peng et al., 2011; using 10 µM CytD). Moreover, in HeLa cells, retrograde transport of virions is sensitive to CytD (2 µM) (Schelhaas et al., 2008). Hence, under CytD, active PsV recruitment should be largely inhibited. Frequently, we observe patches of confluent cells which are common to HaCaT cells. Cells at the center of these patches are dismissed during imaging, because hardly any PsVs are bound to their basal membrane, indicating that PsVs do rather not reach this area by passive diffusion. Instead, we focus on isolated HaCaT cells or cells at the periphery of cell patches. At these cells, we find more PsVs per cell than one would expect from the employed ≍ 50 viral genome equivalents (vge) per cell, indicating that PsVs are unequally distributed between the cells. Moreover, in particular after CytD treatment, these PsVs usually are not homogenously distributed around the cell but rather concentrate at one region (Figures 1 to 4). In later experiments (Figures 5 to 8), we investigate the recruitment of PsVs from these regions, defining ROIs for analysis that cover PsVs at the periphery and the cell body (see Supplementary Figures 5A and 8A).

Earlier version: CytD arrests PsV translocation from the ECM to the cell body.

Revised version: CytD arrests PsV recruitment from the ECM to the cell body.

Earlier version: (A) In the absence (Control) or presence of 10 µg/ml CytD (CytD), HaCaT cells were incubated with PsVs at 37 °C for 5 h. Then, cells were fixed, washed and stained with the cell membrane dye TMA-DPH (gray lookup table (LUT)). PsVs (magenta LUT) were visualized through coupling a dye (6-FAM Azide) to the encapsidated plasmid by click-chemistry. Indirect immunolabeling was employed for staining of HS (AlexaFluor 594; cyan LUT) and Itgα6 (STARRED; green LUT). Imaging was realized with epi-fluorescence microscopy. White lines delineate the main cell body; lines were created with reference to the TMA-DPH membrane staining. (B) Same pre-treatment of cells as in (A), with an additional condition where CytD was added 1 h after the PsVs (CytD after 1 h). Prior to fixation, membrane sheets were generated and F-actin was stained with phalloidin coupled to iFluor488 (green LUT). The capsid protein L1 of the PsVs and CD151 were stained by immunofluorescence using primary antibodies in combination with AlexaFluor 594-labelled (L1, magenta LUT) and STAR RED-labelled (CD151, not shown in this figure for clarity reasons) secondary antibodies. Images of phalloidin and L1 staining were acquired in the confocal and STED mode of a STED microscope, respectively. Arrows in the PsV-L1 images point towards accumulated PsVs that after CytD are more frequently observed than in the control (see text). (C) Analysis of images as shown in (B) using ROIs covering the whole image. PsV maxima were detected and their intensities were quantified in a circular 125 nm diameter region of interest (ROI), followed by background correction. Values are given as means ± SD (n = 3; one biological replicate includes per condition the average of 14 - 15 analyzed membranes (intensity values of one membrane sheet were averaged) with altogether at least 1000 maxima intensity values). (D) PsV maxima intensity distribution of the data in (C). The fraction of PsVs, expressed in percent, is plotted as histogram (10 a.u. bins) against the maxima intensity. Values are given as means ± SD (n = 3). (E) HaCaT cells were treated either for 5 h or 24 h with PsVs, together without (Control) or with 10 µg/ml CytD (CytD). In case of the 5 h incubation, after removal of the PsVs/CytD cells were incubated for another 19 h in medium (in total 24 h). After a total of 24 h incubation, the luciferase activity of lysed cells was measured, yielding the infection rate that was normalized to LDH, resulting in the normalized infection rate. The normalized infection rate was further related to the mean normalized infection rate of the 5 h control, set to 100%, yielding the relative infection rate. Values are given as means ± SD (n = 3 biological replicates; the value of one biological replicate is the average of three technical replicates). (C) and (E). Statistical differences between Control and CytD was analyzed by using the two-tailed, unpaired student’s t-test (n = 3, for details see methods). a.u., arbitrary units.

Revised version: (A) In the absence (Control) or presence of 10 µg/ml CytD (CytD), HaCaT cells were incubated with PsVs at 37 °C for 5 h. Then, cells were fixed, washed and stained with the cell membrane dye TMA-DPH (gray lookup table (LUT)). PsVs (magenta LUT) were visualized through coupling a dye (6-FAM Azide) to the encapsidated plasmid by click-chemistry. Indirect immunolabeling was employed for staining of HS (AlexaFluor 594; cyan LUT) and Itgα6 (STARRED; green LUT). Imaging was realized with epi-fluorescence microscopy. White lines delineate the main cell body; lines were created with reference to the TMA-DPH membrane staining. (B) Same pre-treatment of cells as in (A), with an additional condition where CytD was added 1 h after the PsVs (CytD after 1 h). Prior to fixation, membrane sheets were generated and F-actin was stained with phalloidin coupled to iFluor488 (green LUT). The capsid protein L1 of the PsVs and CD151 were stained by immunofluorescence using primary antibodies in combination with AlexaFluor 594-labelled (L1, magenta LUT) and STAR RED-labelled (CD151, not shown in this figure for clarity reasons) secondary antibodies. Images of phalloidin and L1 staining were acquired in the confocal and STED mode of a STED microscope, respectively. Arrows in the PsV-L1 images point towards accumulated PsVs that after CytD are more frequently observed than in the control (see text). (C) Analysis of images as shown in (B) using ROIs covering the whole image. PsV maxima were detected and their intensities were quantified in a circular 125 nm diameter region of interest (ROI), followed by background correction. Values are given as means ± SD (n = 3; one biological replicate includes per condition the average of 14 - 15 analyzed membranes (intensity values of one membrane sheet were averaged) with altogether at least 1000 maxima intensity values). (D) PsV maxima intensity distribution of the data in (C). The fraction of PsVs, expressed in percent, is plotted as histogram (10 a.u. bins) against the maxima intensity. Values are given as means ± SD (n = 3). (E) HaCaT cells were treated either for 5 h or 24 h with PsVs, without (Control) or with 10 µg/ml CytD (CytD). In case of the 5 h incubation, after removal of the PsVs/CytD cells were incubated for another 19 h in medium (in total 24 h). After a total of 24 h incubation, the luciferase activity of lysed cells was measured, yielding the infection rate that was normalized to LDH, resulting in the normalized infection rate. The normalized infection rate was further related to the mean normalized infection rate of the 5 h control, set to 100%, yielding the relative infection rate. Values are given as means ± SD (n = 3 biological replicates; the value of one biological replicate is the average of three technical replicates). (C) and (E). Statistical differences between Control and CytD was analyzed by using the two-tailed, unpaired student’s t-test (n = 3, for details see methods). a.u., arbitrary units.

Added in the revised version: In Figure 1, after the 5 h incubation, cells were fixed and stained by a membrane marker, the dye TMA-DPH, to mark the main cell body (Figure 1A, gray lookup table (LUT)). Additionally, PsVs were visualized by click-chemistry (magenta LUT), and HS (cyan LUT) and the membrane protein Itgα6 (green LUT) were stained by antibodies. In order to monitor only cell surface events, we stain the cells without prior cell permeabilization (still, fixation perforates the cell membrane to some extent). In the control, PsVs locate strongly in the cell body area (Figure 1A, upper row; see white outline that is based on TMA-DPH staining). In contrast, when CytD was present, the cell body area is largely devoid of PsVs (Figure 1A, lower row; please note that the large central bright area is mainly caused by autofluorescence). Instead, PsVs often accumulate adjacent to the cell body. This adjacent region, that is up to several µm wide and rich in HS, likely includes the ECM. We conclude that upon inhibition of actin-dynamics PsVs do not efficiently reach the basal cell membrane but rather remain accumulated in the ECM area. In contrast, in the control they are actively recruited to the basal cell membrane. Please note that under all conditions PsVs will bind to receptors located outside of the imaged basal membrane, e.g., the entire upper cell membrane, but these PsVs are not visible in the micrographs and therefore not included in the analysis.

Earlier version: Comparing in Figure 1A the PsV brightness in the upper (control, magenta LUT) and lower (CytD, magenta LUT) panels suggests that the amount of PsVs does not differ greatly. This is surprising, as CytD is not reported to inhibit enzymes involved in capsid modification or HS/HSPGs processing, after which the PsVs supposedly leave the ECM. If this were the case, one would expect no difference in the number of PsVs at the periphery between cells that were treated with or without CytD for 5 h. Therefore, we conclude that PsVs remain ECM associated while ready for binding to a secondary receptor. However, a quantification of the PsVs with and without CytD is not possible in this experiment, due to interference from autofluorescence (Figure 1A, large magenta area in the center). To study PsVs in the absence of autofluorescence, we employed unroofed cells (Heuser, 2000), also referred to as membrane sheets. After pre-treatment of cells as above, membrane sheets are generated by brief ultrasound pulses that remove the upper cellular parts (and thus any intracellular autofluorescence), leaving behind the extracellular matrix and the basal membrane along with the bound PsVs. During imaging, membrane sheets are identified via staining of F-actin with fluorescently labelled phalloidin (Figure 1B, green LUT) that marks the main cell body cortex together with lamellipodia and filopodia. In this experiment we observe that PsVs, not always but frequently, accumulate adjacent to the isolated basal membrane (defined by F-actin staining, Figure 1B, green LUT) after CytD treatment, which differs from the control (Figure 1B, magenta LUT).

Revised version: Comparing in Figure 1A the PsV brightness in the upper (control, magenta LUT) and lower (CytD, magenta LUT) panel suggests that the PsV amounts do not differ greatly. This is surprising, as CytD is not reported to inhibit enzymes involved in capsid modification or HS/HSPGs processing, after which the PsVs supposedly leave the ECM. Provided PsVs processed under CytD would leave the ECM, no major difference in the number of PsVs at the periphery between cells that were treated with or without CytD for 5 h is expected. Therefore, it appears that primed PsVs remain ECM associated. Due to interference from autofluorescence (Figure 1A, large magenta areas in the centers), a quantification of the PsVs with and without CytD is not possible. To study PsVs in the absence of autofluorescence, we employed unroofed cells (Heuser, 2000), also referred to as membrane sheets. After pre-treatment of cells as above, membrane sheets are generated by brief ultrasound pulses that remove the upper cellular parts (and thus any intracellular autofluorescence), leaving behind the extracellular matrix and the basal membrane along with the bound PsVs. During imaging, membrane sheets are identified via staining of F-actin with fluorescently labelled phalloidin (Figure 1B, green LUT) that marks the main cell body cortex together with lamellipodia and filopodia. In this experiment we again observe that PsVs, not always but more frequently, after CytD treatment accumulate adjacent to the isolated basal membrane (Figure 1B, magenta LUT, see arrows) that is defined by F-actin staining (Figure 1B, green LUT).

Earlier version: The PsV maxima density (Figure 1B, magenta LUT, detected as local maxima of antibody stained L1) and the PsV intensity are quantified in an area covering the cell body and the periphery. Compared to the control, CytD reduces the PsV maxima density by 26% (the PsV maxima density in the control and CytD treated cells is 0.7 and 0.52 PsVs/µm2, respectively), whereas the maxima intensity increases by 41% (Figure 1C). A histogram of the PsV maxima intensities illustrates that CytD broadens the intensity distribution towards brighter PsVs (Figure 1D). The increase of the PsV intensity after CytD treatment can also be appreciated comparing the upper and middle/lower L1 image of Figure 1B (shown at the same settings of brightness and contrast). As CytD is unlikely to act directly on PsVs (e.g., by fusing them to larger and brighter particles), we assume that the underlying cause of brighter maxima is the resolution limit of the microscope. The densely accumulated PsVs are no longer resolved and merge to brighter spots. These poorly resolved maxima are on average 41% brighter but of them there are 26% fewer. The decrease in maxima density is in the order of magnitude of the increase in intensity, which yields a roughly similar total signal in both conditions.

Revised version: The PsV maxima density (Figure 1B, magenta LUT, detected as local maxima of antibody stained L1) and the PsV intensity are quantified in an area covering the cell body and the periphery. Compared to the control, CytD reduces the PsV maxima density by 26% (the PsV maxima density in the control and CytD treated cells is 0.7 and 0.52 PsVs/µm2, respectively), whereas the maxima intensity increases by 41% (Figure 1C). A histogram of the PsV maxima intensities illustrates that CytD broadens the intensity distribution towards brighter PsVs (Figure 1D, compare black and gray trace). The increase of the PsV intensity after CytD treatment can also be appreciated comparing the upper and middle/lower L1 images in Figure 1B (shown at the same settings of brightness and contrast). As CytD is unlikely to act directly on PsVs (e.g., by fusing them to larger and brighter particles), we assume that the underlying cause of brighter maxima is the resolution limit of the microscope. The densely accumulated PsVs are no longer resolved, overlap, and by this merge to brighter spots. These poorly resolved maxima are on average 41% brighter but of them there are 26% fewer. The decrease in maxima density is in the order of magnitude of the increase in intensity, which yields a roughly similar total signal in both conditions. Hence, active recruitment apparently thins out the PsV accumulations by translocating PsVs from the ECM to the cell body.

Earlier version: We addressed whether PsVs would still be accumulated when adding CytD 1 h after the PsVs rather than simultaneously (Figure 1B, lower panel, CytD after 1 h). As in the CytD condition, we frequently see accumulated PsVs adjacent to the isolated basal membrane (Figure 1B, arrows point towards PsV accumulations). The PsV maxima intensity distribution is broader when compared to the control but narrower than with CytD present the entire time (Figure 1D). We conclude that complete PsV binding and translocation takes longer than 1 h, as otherwise control and CytD after 1 h should be more similar.

Revised version: We addressed whether PsVs would still be accumulated when adding CytD 1 h after the PsVs, rather than adding them simultaneously (Figure 1B, lower panel, CytD after 1 h). As in the CytD condition, accumulated PsVs adjacent to the basal membrane are present more frequently than in the control (Figure 1B, arrows). The PsV maxima intensity distribution (Fig. 1D, green trace) is slightly broader when compared to the control, but narrower than with CytD present for the entire time (Figure 1D), suggesting that PsV binding and recruitment take more than 1 h.

Earlier version: At this point, one could conclude from the data that CytD arrests PsVs on their journey towards the cell body and that PsVs are not collected from the cells, but remain trapped in the ECM. How efficient PsV trapping is cannot be determined exactly for several reasons. First, our assay is not precise for technical reasons, as the above-mentioned limited resolution results in an underestimation of the number of PsVs the stronger they are accumulated. Second, we do not know how many PsVs in the control become endocytosed during the 5 h of incubation. This reduces the PsVs in the control, which is our value for comparison. Moreover, it should be noted that both the degree of PsV accumulation as well as the number of PsVs are highly variable. Even in the absence of CytD, we occasionally find cells with accumulated PsVs at the periphery, whereas such cells are more frequently observed after CytD. This causes a large variability within the data. Still, we consider larger PsV accumulations after CytD as a typical effect (please note that for illustration of this effect we show in the Figures such examples). Despite these limitations, it appears that within the 5 h incubation period with CytD a large fraction of PsVs remains associated with the ECM.

Revised version: Figure 1 indicates that under CytD, when PsVs are not actively collected from the cells, they remain trapped in the ECM next to the cell body. However, for several reasons it cannot be determined exactly how efficient the preservation of accumulated PsVs is. First, our assay is not precise for technical reasons, as the above-mentioned limited resolution results in an underestimation of the number of PsVs the stronger they are accumulated. Second, we do not know how many PsVs in the control are endocytosed during the 5 h of incubation, which is an important number, as adding theses PsVs to the ones being present would be the adequate reference value for comparison to CytD. Moreover, it should be noted that both the degree of PsV accumulation as well as the number of PsVs are highly variable. Even in the control, we occasionally find cells with accumulated PsVs at the periphery, while after CytD we not always encounter accumulated PsVs. This causes a large variability within the data. Still, we consider larger PsV accumulations as a typical effect of CytD incubation (please note that for illustration of this effect we show in the Figures such examples). In any case, we conclude that within the 5 h incubation period with CytD a large fraction of ECM associated PsVs is not able to reach the cell body as in the control, but remains located in its periphery.

Removed in the revised version: During the CytD incubation, PsVs bind to HSPGs of the basolateral membrane for 5 h. Still, in the cell body area hardly any PsVs are present (0.14 PsV/µm2, Supplementary Figure 1B). In the control, the PsV density is several-fold larger (Supplementary Figure 1B). This is expected, as the PsVs bind to the ECM and translocate to the cell body. We wondered whether there are more binding sites at the basal membrane that, however, remain inaccessible to PsVs by diffusion, because of the insufficient space between glass-coverslip and basolateral membrane. For clarification, we incubated EDTA detached HaCaT cells in suspension with PsVs for 1 h at 4 °C, followed by re-attachment for 1 h. Under these conditions, we find a PsV density 12.4-fold larger than after 5 h of CytD incubation of adhered cells (Supplementary Figure 1B and D). However, it should be noted that these values cannot be directly compared. Aside from the different treatments, another difference lies in the size of the basal membrane, as re-attachment of cells is not complete after only 1 h (compare size of adhered membranes in Supplementary Figure 1A and C). Therefore, the imaged membranes are likely strongly ruffled, which results in the underestimation of the size of the adhered membrane. As a result, we overestimate the PsVs per µm2 (please note that we cannot re-attach cells for longer times as we would then lose PsVs due to endocytosis). On the other hand, we would underestimate the PsV density at the basal membrane if after re-attachment we image in part also some apical membrane. In any case, the experiment suggests that PsVs bind more efficiently if membrane surface receptors are accessible by diffusion. This is in support of the above notion that the basal membrane may provide more entry receptors than one would expect from the low density of PsVs bound after 5 h CytD (Supplementary Figure 1B). This suggests that under our assay conditions, PsVs cannot easily bypass the translocation from the ECM to the cell body by diffusing directly to the basal membrane. Hence, the large majority of PsVs that enter the cell were previously bound to the ECM. Therefore, HaCaT cells serve as an ideal model for studying the transfer of ECM bound HPV particles to the cell surface, which is similar to in vivo infection of basal keratinocytes after binding to the basement membrane (Bienkowska-Haba et al., 2018; Day and Schelhaas, 2014; Kines et al., 2009; Schiller et al., 2010).

The blocking of PsV translocation by cytochalasin D is reversible

Earlier version: CytD not only arrests the transport of PsVs to the cell but is known to block other actin-dependent processes, which strongly affects the physiology of the cell. Therefore, we investigated whether PsVs would proceed normally on their infection pathway once CytD is washed off. In order to allow for recovery from CytD, cells were washed after the 5 h simultaneous PsV/CytD treatment, then incubated for another 19 h and the infection rate was subsequently determined by measuring the luciferase activity in the cell lysate. CytD reduced luciferase activity by 29%, as compared to the control (Figure 1E; left). A certain degree of reduced infectivity is to be expected, as PsVs have ≈20% less time to complete infection as compared to cells that were not treated with CytD. When PsVs were not washed off after 5 h but left for the entire 24 h incubation, the infection rate increased by 52%, whereas continuous treatment with CytD results in no infection (Figure 1E; right). The latter is in line with previous studies (Schelhaas et al., 2008; Selinka et al., 2002; Selinka et al., 2007; Spoden et al., 2013) showing that CytD is a strong inhibitor of HPV infection. In any case, PsVs proceed on the infection pathway upon removal of CytD. The reduced infection rate of 29% can be explained by the 5 h blocking at the beginning of the incubation period that should delay the time course of infection. Therefore, we propose that CytD is suitable for transiently arresting PsVs in a state between primary attachment to HS and cell body secondary receptor binding.

Revised version: CytD not only arrests the recruitment of PsVs to the cell body but is known to block other actin-dependent processes, which strongly affects the physiology of the cell. Therefore, we investigated whether PsVs would proceed normally on their infection pathway once CytD is washed off. In order to allow for recovery from CytD, cells were washed after the 5 h PsV/CytD treatment, then incubated for another 19 h and the infection rate was determined by measuring the luciferase activity in the cell lysate. CytD reduced luciferase activity by 29%, as compared to control cells incubated only with PsVs (Figure 1E; left). When PsVs were not washed off after 5 h but left for the entire 24 h incubation, the infection rate increased by 52%, whereas continuous treatment with CytD blocks infection (Figure 1E; right). The latter is in line with previous studies (Schelhaas et al., 2008; Selinka et al., 2002; Selinka et al., 2007; Spoden et al., 2013) showing that CytD is a strong inhibitor of HPV infection. In any case, PsVs apparently are able to proceed on the infection pathway upon removal of CytD. The reduced infection rate of 29% can be explained by the 5 h blocking at the beginning of the incubation period that should delay the time course of infection. Altogether, we propose that CytD is suitable for transiently arresting PsVs in a state between primary attachment to HS and cell body receptor binding.

Earlier version: Next, we investigated the onset of PsV translocation after CytD wash off. Cells were treated with PsVs and CytD for 5 h, washed, and incubated further without PsVs/CytD for 0 min, 30 min or 60 min, followed by fixation and staining for PsVs (Figure 2A, magenta LUT) and F-actin (Figure 2A, green LUT). The antibody staining against PsVs, like above the click-chemistry staining, is again highly variable. For analyzing the translocation process, we calculated the Pearson correlation coefficient (PCC) between PsV-L1 and F-actin staining. The PCC quantifies the similarity between two variables, in this case the pixel values of the two images. The PCC results in the value of 1 if images overlap perfectly and -1 for an image and its negative. After 5 h of CytD, PsVs again accumulate at the edge of the F-actin staining (Figure 2A, compare to Figure 1B that shows membrane sheets). The distal PsVs and the F-actin staining partially exclude each other, thus at 0 min we obtain a negative PCC (Figure 2A and D). After 30 min and 60 min, PsVs mostly overlap with F-actin stained areas, which results in positive PCCs (Figure 2A and D). The time course of the change in PCC values suggests an onset of translocation recovery within 30 min after CytD removal.

Revised version: Next, we investigated the onset of active PsV recruitment after CytD wash off. Cells were treated with PsVs and CytD for 5 h, washed, and incubated further without PsVs/CytD for 0 min, 30 min or 60 min, followed by fixation and staining for PsVs (Figure 2A, magenta LUT) and F-actin (Figure 2A, green LUT). The antibody staining against PsVs, like above the click-chemistry staining, is again highly variable. As a measure of recruitment to the cell body, we calculated the Pearson correlation coefficient (PCC) between PsV-L1 and F-actin staining. The PCC quantifies the similarity between two variables, in this case the pixel values of the two images. The PCC results in the value of 1 if images are identical and -1 for an image and its negative. After 5 h of CytD, PsVs often are found accumulated, as seen before, at the edge of the F-actin staining (Figure 2A). The distal PsVs and the F-actin staining partially exclude each other, thus at 0 min we obtain a negative PCC (Figure 2A and D). After 30 min and 60 min, PsVs mostly overlap with F-actin stained areas, which results in positive PCCs (Figure 2A and D). The time course of the change in PCC values suggests an onset of recruitment after CytD removal within 30 min.

Earlier version: Transfer of PsVs from the ECM to the cell body requires PsV priming.

Revised version: Recruitment of PsVs from the ECM to the cell body requires PsV priming.

HaCaT cells were pre-incubated for 5 h at 37 °C with PsVs, in the presence of (A) 10 µg/ml CytD (CytD), (B) 10 µg/ml CytD and 100 µM leupeptin (CytD/leupeptin), or (C) 10 µg/ml CytD and 5 µM Furin inhibitor I (CytD/Furin inhibitor I). Afterwards, cells were washed and incubated without PsVs/inhibitors further for 0 min, 30 min or 60 min, before they were fixed and stained by indirect immunofluorescence for L1 (STAR GREEN, magenta LUT) and for F-actin by iFluor647-labelled phalloidin (green LUT). PsVs-L1 and F-actin staining were imaged in the confocal mode of a STED microscope. (D) For determination of the Pearson correlation coefficient (PCC) between PsV-L1 (magenta LUT) and Phalloidin (green LUT), we placed large ROIs onto the images that covered mainly the cell body but included parts of the cell periphery as well. The PCC was plotted over time. Values are given as means ± SD (n = 3 biological replicates). Statistical difference between CytD and CytD/inhibitors was analyzed by using the two-tailed, unpaired student’s t-test (n = 3, for details see methods).

Earlier version: Assessing previously proposed steps in HPV infection in our assay

Revised version: Confirming in our assay previously proposed steps in HPV infection

Earlier version: PsVs reportedly bind tightly to their primary attachment site until the capsid surface undergoes structural changes, which involves modifications of L1 and L2 (see above). Consequently, we would expect that inhibition of L1 processing during the CytD incubation prevents the recovery of PsV translocation from the ECM to the cell body (Figure 2A and D). To test for this possibility, as employed in earlier studies, the protease inhibitor leupeptin was used to inhibit proteases including KLK8 which is required for L1 cleavage (Cerqueira et al., 2015). Employing this inhibitor, the PCC between PsV-L1 and F-actin staining remains negative after CytD removal, showing that for translocation indeed the action of proteases is required (Figure 2B and D). In contrast, inhibition of L2 cleavage by a furin specific inhibitor has no effect on the PCC (Figure 2C and D). However, it should be noted that we occasionally observe PsVs not completely translocating but accumulating at the border of the F-actin stained area (for example see Figure 2C (60 min)). This results in an increase of the PCC almost equal to complete translocation, explaining why the PCC remains unaffected despite a furin inhibitory effect. Hence, furin inhibition may have some effect on translocation that, however, is undetected in this type of analysis.

Revised version: As outlined above, PsVs bind tightly to their primary attachment site until the capsid surface undergoes structural changes, which involves modifications of L1 and L2. Consequently, we expect that inhibition of L1 processing during the CytD incubation should inhibit PsV recruitment after CytD removal (Figure 2A and D). To test for this possibility, as employed in earlier studies, the protease inhibitor leupeptin was used to inhibit proteases including KLK8 which is required for L1 cleavage (Cerqueira et al., 2015). Employing this inhibitor, the PCC between PsV-L1 and F-actin staining remains negative after CytD removal, showing that for active recruitment indeed the prior action of proteases is a prerequisite (Figure 2B and D). Moreover, the experiment suggests that without PsV priming the PCC between PsV-L1 and F-actin does not increase, for instance, due to cell spreading after CytD removal. In contrast, inhibition of L2 cleavage by a furin specific inhibitor has no effect on the PCC (Figure 2C and D). However, it should be noted that we occasionally observe PsVs not undergoing complete translocation. Instead, they remain accumulated at the border of the F-actin stained area (for example see Figure 2C, 60 min). This results in an increase of the PCC as in complete translocation, explaining why the PCC changes like in the control, despite of a furin effect. Hence, furin may have some effect on a later recruitment step that, however, is undetected in this type of analysis.

Earlier version: As outlined above, during the 5 h incubation with CytD, proteases in the ECM are expected to cleave HS chains. These cleavage products should be able to diffuse out of the ECM, unless they remain associated with non-translocating PsVs. On the other hand, in the control, less HS cleavage products are trapped and if trapped the PsV leave the ECM. So far, and throughout the entire study, we stain the cells without prior cell permeabilization (albeit fixation perforates the cell membrane to some extent) in order to monitor only cell surface events. Using an antibody that reacts with an epitope in native heparan sulfate chains, only after CytD and if PsVs were present, the level of HS staining was significantly increased (Figure 3B). As shown in Figure 3A, stronger HS staining at PsVs (open arrows) and as well in PsV free areas (closed arrows) was observed.

Revised version: As outlined above, during the 5 h incubation with CytD, proteases in the ECM are expected to cleave HS chains. These cleavage products should be able to diffuse out of the ECM, unless they remain associated with non-translocating PsVs, present in particular under CytD. Using an antibody that reacts with an epitope present in native heparan sulfate chains, we find that only after CytD and if PsVs were present, the level of HS staining is significantly increased (Figure 3B). As shown in Figure 3A, this stronger HS staining is observed in areas with PsVs (open arrows) and as well in PsV free areas (closed arrows).

Increased HS intensity after incubation with PsVs and CytD.

Earlier version: (A) HaCaT cells were incubated without (top) and with (bottom) PsVs at 37 °C for 5 h, in the absence (Control, upper panels) or presence of 10 µg/ml CytD (CytD, lower panels). Afterwards, cells were washed, fixed and stained. Immunofluorescence was used for L1 (STAR GREEN; magenta LUT) and for HS (AlexaFluor 594; cyan LUT). F-actin was stained by iFluor647-labelled phalloidin (green LUT). PsVs-L1, HS and F-actin staining were imaged in the confocal mode of a STED microscope. The open arrow marks a region where PsVs overlap with HS. The closed arrow marks a region devoid of PsVs showing strong HS staining. (B) For analysis of the mean HS intensity, we placed large ROIs onto the images covering mainly the cell body but including parts of the cell periphery as well. Values are given as means ± SD (n = 3 biological replicates). Statistical differences between Control and CytD were analyzed by using the two-tailed, unpaired student’s t-test (n = 3, for details see methods).

Revised version: (A) HaCaT cells were incubated without (top) and with (bottom) PsVs at 37 °C for 5 h, in the absence (Control, upper panels) or presence of 10 µg/ml CytD (CytD, lower panels). Afterwards, cells were washed, fixed and stained. Immunofluorescence was used for L1 (STAR GREEN; magenta LUT) and for HS (AlexaFluor 594; cyan LUT) staining. F-actin was stained by iFluor647-labelled phalloidin (green LUT). PsVs-L1, HS and F-actin staining were imaged in the confocal mode of a STED microscope. The open arrow marks a region where PsVs overlap with HS. The closed arrow marks a region devoid of PsVs showing strong HS staining. (B) For analysis of the mean HS intensity, we placed large ROIs onto the images covering mainly the cell body but including parts of the cell periphery as well. Values are given as means ± SD (n = 3 biological replicates). Statistical differences between Control and CytD were analyzed by using the two-tailed, unpaired student’s t-test (n = 3, for details see methods).

Earlier version: As additional control and shown in Supplementary Figure 2, we used an antibody that reacts with a HS neo-epitope generated by heparitinase-treated heparan sulfate chains (Yokoyama et al., 1999; for details see methods). This neo-epitope staining is independent of the presence of CytD and the incubation time, suggesting that CytD does not directly affect HS processing. Collectively, our findings indicate that without actin-dependent PsV translocation HS cleavage products are retained in the ECM, consistent with the hypothesis that cleaved HS remains associated with PsVs (Ozbun and Campos, 2021).

Revised version: Next, we used an antibody that reacts with a HS neo-epitope generated by heparitinase-treated heparan sulfate chains (Yokoyama et al., 1999; for details see methods). This neo-epitope staining is not affected by CytD and the incubation time (Supplementary Figure 1), suggesting that CytD does not directly affect HS processing. Collectively, our findings indicate that without actin-dependent PsV recruitment HS cleavage products are retained in the ECM, consistent with the hypothesis that cleaved HS remains associated with PsVs (Ozbun and Campos, 2021).

Earlier version: The transfer of PsVs from the ECM to the cell body is fast and efficient

Revised version: Fast recruitment of PsVs to the cell body

Earlier version: Next, we studied the time course of PsV translocation in more detail. Cells were treated as in Figure 1A (in Figure 4, the 0 min time point is identical), followed by removal of PsVs/CytD. Different from Figure 2, we did not calculate the PCC that provides only an approximate idea of the time course of translocation, but monitored the diminishment of the integrated density of PsVs locating at the cell periphery after 0, 15, 30 and 60 min. The cell periphery is delineated from the cell body by the TMA-DPH membrane staining (Figure 4A, the cell periphery is defined by the area enclosed by the two white lines in the images shown with the magenta LUT).

Revised version: Next, we studied the time course of PsV recruitment at higher time resolution. Cells were treated as in Figure 1A (in Figure 4, the 0 min time point is identical), followed by removal of PsVs/CytD. Different from Figure 2, we did not calculate the PCC that provides only an approximate idea of the time course, but monitored the diminishment of the integrated density of PsVs locating at the cell periphery, after 0, 15, 30 and 60 min. To delineate the cell periphery from the cell body, the TMA-DPH membrane staining is used as a reference (Figure 4A, the cell periphery is defined by the area enclosed by the two white lines in the PsV-DNA images (magenta LUT)).

Fast diminishment of accumulated PsVs at the cell periphery after removal of CytD.

Earlier version: (A) HaCaT cells were incubated with PsVs at 37 °C for 5 h, in the absence (Control) or presence of 10 µg/ml CytD (CytD). Then, cells were washed and incubated for the indicated time periods without PsVs/CytD, before they were fixed and stained as in Figure 1 (t = 0 min is identical to Figure 1A; for clarity we show only the membrane (gray LUT; images are shown at different settings of brightness and contrast) and the PsV-DNA staining (magenta LUT; images are shown at the same settings of brightness and contrast). The white lines in the membrane images delineate the cell body from the periphery. They were created with reference to the membrane staining (for details see methods). Using the cell body delineation as starting point, an up to 30-pixel broad area was created (PsV channel, magenta LUT; see area enclosed by the smoother white lines and the cell body delineation lines). The area enclosed by the two white lines defines the cell periphery. (B) The PsV-DNA signal of the periphery was quantified as integrated density, background corrected, and plotted over time. Values are given as means ± SD (n = 3 biological replicates). The statistical difference between the same time points of Control and CytD were analyzed by using the two-tailed, unpaired student’s t-test (n = 3, for details see methods).

Revised version: (A) HaCaT cells were incubated with PsVs at 37 °C for 5 h, in the absence (Control) or presence of 10 µg/ml CytD (CytD). Then, cells were washed and incubated for the indicated time periods without PsVs/CytD, before they were fixed and stained as in Figure 1 (t = 0 min is identical to Figure 1A; for clarity we show only the membrane (gray LUT; images are shown at different settings of brightness and contrast) and the PsV-DNA staining (magenta LUT; images are shown at the same settings of brightness and contrast). The white lines in the membrane images delineate the cell body from the periphery. They were created with reference to the membrane staining (for details see methods). Using the cell body delineation as starting point, an up to 30-pixel broad area was created (PsV channel, magenta LUT; areas enclosed by the smoother white lines and the cell body delineation lines). The areas enclosed by the white lines define the cell peripheries. (B) The PsV-DNA signal of the periphery was quantified as integrated density, background corrected, and plotted over time. Values are given as means ± SD (n = 3 biological replicates). The statistical difference between the same time points of Control and CytD were analyzed by using the two-tailed, unpaired student’s t-test (n = 3, for details see methods).

Earlier version: At 0 min, compared to the control, CytD causes a 6-fold increase of PsV signal in the periphery (Figure 4B). This increase is more than halved when cells were incubated for 15 min, and after 30 min, the level of the control is reached (Figure 4B). This suggests that the half-time of PsV translocation from the periphery to the cell body is about 15 min. In fact, the half-time maybe longer, as we cannot exclude that cell spreading after CytD removal contributes to less PsVs measured in the cell periphery. In any case, translocation is fast and efficient and not a major component in the time course of infection. We conclude that translocation is not responsible for the asynchronous virion uptake observed for HPV PsVs.

Revised version: At 0 min, compared to the control, CytD causes in the periphery a 6-fold increase of PsV signal (Figure 4B). This increase is more than halved when cells were incubated for 15 min, and after 30 min, the level of the control is reached (Figure 4B). This suggests that the half-time of PsV recruitment from the periphery to the cell body is about 15 min. Hence, active recruitment is fast and therefore cannot be a bottleneck in the time course of infection. We conclude that active recruitment is not responsible for the asynchronous virion uptake observed for HPV PsVs.

Earlier version: Translocation of PsVs to CD151

Revised version: Recruitment of PsVs to CD151

Earlier version: During translocation, at some point PsVs may approach to cell surface HPV entry factors, like the tetraspanin CD151. For studying this possibility, we employ superresolution STED microscopy, analyzing the association of PsVs with CD151 over time. Cells are treated and monitored as above, but with an extended time window of up to 180 min, as cell surface processes are expected to take longer than an hour.

Revised version: PsVs may approach the entry factor CD151 early, already during recruitment. For studying this possibility, we employ superresolution STED microscopy, analyzing the association of PsVs with CD151 over time. Cells are treated and monitored as above, but with an extended time window of up to 180 min, as cell surface processes are expected to take longer than an hour. After fixation, PsVs and CD151 were double-stained with antibodies against L1 (Figure 5A, magenta LUT) and CD151 (Figure 5A, green LUT). In addition, we stained F-actin with fluorescent labelled phalloidin (shown in Supplementary Figures 2 and 3). We simultaneously image L1 and F-actin in the confocal and CD151 in the STED channel.

Removed in the revised version: After fixation, PsVs and CD151 were double-stained with antibodies against L1 (Figure 5A, magenta LUT) and CD151 (Figure 5A, green LUT). In addition, we stained F-actin with fluorescent labelled phalloidin. We simultaneously image L1 and F-actin in the confocal and CD151 in the STED channel.

Earlier version: Association between PsVs and CD151 is an early event in the infection cascade.

Revised version: Association between PsVs and CD151 occurs early in the infection cascade.

Earlier version: (A) HaCaT cells were incubated with PsVs at 37 °C for 5 h, in the absence (Control, upper panels) or presence of 10 µg/ml CytD (CytD, lower panels). Afterwards, cells were washed and incubated without PsVs/CytD further for 0 min, 30 min, 60 min or 180 min, before they were fixed and stained by indirect immunofluorescence for L1 (STAR GREEN; magenta LUT) and for CD151 (AlexaFluor 594; green LUT), and for F-actin by iFluor647-labelled phalloidin (here not shown for clarity, please see Supplementary Figure 3 for F-actin illustrating the variability of filopodia after CytD treatment). The bottom rows show magnified views of the merged images from the regions marked by the white boxes. PsV-L1 and F-actin staining were imaged in the confocal and CD151 staining in the STED mode of a STED microscope, respectively. Therefore, compared to CD151, the PsVs are less resolved and also appear much larger than their real physical size (see magnified views). CytD/180 min, arrows mark presumably endocytic structures that formed in the central cell body region (for more examples see Supplementary Figure 5). For analysis, we placed large ROIs onto the images that covered mainly the cell body but included parts of the cell periphery as well (for an example ROI see Supplementary Figure 6A). (B) The mean CD151 intensity was measured and plotted over time. (C) The PCC between PsV-L1 (magenta LUT) and CD151 (green LUT) was calculated and plotted over time. (D) The fraction of PsVs (in percent) that have a distance to the next neighbored CD151 maximum ≤ 80 nm, which we define as closely associated, is plotted over time. Please note that the values in (D) were corrected for random background association (for details see Supplementary Figure 7). Two examples of PsVs (each marked by an asterisk) from the CytD/30 min (left) and CytD/60 min (right) conditions are shown. The value in the upper left states the shortest distance between the PsV and the next nearest CD151 maximum (marked by an arrow) in nm. Values are given as means ± SD (n = 3 biological replicates). Statistical differences between the same time points of Control and CytD were analyzed by using the two-tailed, unpaired student’s t-test (n = 3, for details see methods).

Revised version: (A) HaCaT cells were incubated with PsVs at 37 °C for 5 h, in the absence (Control, upper panels) or presence of 10 µg/ml CytD (CytD, lower panels). Afterwards, cells were washed and incubated without PsVs/CytD further for 0 min, 30 min, 60 min or 180 min, before they were fixed and stained by indirect immunofluorescence for L1 (STAR GREEN; magenta LUT) and for CD151 (AlexaFluor 594; green LUT), and for F-actin by iFluor647-labelled phalloidin (here not shown for clarity, please see Supplementary Figure 2A for F-actin staining illustrating the variability of filopodia after CytD treatment). The bottom rows show magnified views of the merged images from the regions marked by the white boxes. PsV-L1 and F-actin staining were imaged in the confocal and CD151 staining in the STED mode of a STED microscope, respectively. Therefore, compared to CD151, the PsVs are less resolved and also appear much larger than their real physical size (see magnified views). CytD/180 min, arrows mark presumably endocytic structures that formed in the central cell body region (for more examples see Supplementary Figure 4). For analysis, we placed large ROIs onto the images that covered mainly the cell body but included parts of the cell periphery as well (for an example ROI see Supplementary Figure 5A). (B) The mean CD151 intensity was measured and plotted over time. (C) The PCC between PsV-L1 (magenta LUT) and CD151 (green LUT) was calculated and plotted over time. (D) The fraction of PsVs (in percent) that have a distance to the next neighbored CD151 maximum ≤ 80 nm, which we define as closely associated, is plotted over time. Please note that the values in (D) were corrected for random background association (for details see Supplementary Figure 7). Two examples of PsVs (each marked by an asterisk) from the CytD/30 min (left) and CytD/60 min (right) conditions are shown. The value in the upper left states the shortest distance between the PsV and the next nearest CD151 maximum (marked by an arrow) in nm. Values are given as means ± SD (n = 3 biological replicates). Statistical differences between the same time points of Control and CytD were analyzed by using the two-tailed, unpaired student’s t-test (n = 3, for details see methods).

Earlier version: As shown in Figure 5A, CD151 concentrates in spots scattered across the cell surface. It is also present at cell protrusions that are rich in F-actin and vary strongly in number and shape (Supplementary Figure 3A shows some examples of these filopodia). Compared to the control, CytD again preserves the accumulated PsVs at the cell periphery, and at early time points PsVs appear brighter (Figure 5A; for overview images and F-actin staining see Supplementary Figure 4A and B; for variability of PsV accumulations after CytD see Supplementary Figure 4B).

Revised version: As shown in Figure 5A, CD151 concentrates in spots scattered across the cell surface. It is also present at cell protrusions that are rich in F-actin, and that vary strongly in number and shape (Supplementary Figure 2A shows some examples of these filopodia). Compared to the control, CytD increases the occurrence of accumulated PsVs at the cell periphery, and at early time points PsVs appear brighter (Figure 5A; for overview images and F-actin staining of 0 min/CytD see Supplementary Figure 3A and B; for variability of PsV accumulations after CytD see Supplementary Figure 3B), as already observed in Figure 1B to D.

Earlier version: Initially, we wondered whether, after reaching the cell body, PsVs are transported quickly further towards the center of the cell body. Based on the CD151 image, a cell border region was defined (for details see Supplementary Figure 3B). We counted the number of PsVs in this region, and expressed it as percentage of all PsVs in the image. In the control, the fraction of PsVs in the cell border region was rather low, ranging from 13.4% (0 min) to 7.6% (180 min). After CytD, at 0 min, compared to the control the fraction almost tripled (36.5%). It diminished to 16.9% after 60 min, and 12.3% after 180 min (Supplementary Figure 3C). This suggests that PsVs only briefly remain in the cell periphery.

Revised version: Initially, we wondered whether, after reaching the cell body edge, PsVs are transported quickly further towards the center. Based on the CD151 image, a cell border region was defined (for details see Supplementary Figure 2B). We counted the number of PsVs in this region, and expressed it as percentage of all PsVs in the image. In the control, the fraction of PsVs in the cell border region was rather low, ranging from 13.4% (0 min) to 7.6% (180 min). After CytD, at 0 min, compared to the control the fraction almost tripled (36.5%). It diminished to 16.9% after 60 min, and 12.3% after 180 min (Supplementary Figure 2C). This suggests that most PsVs leave the cell periphery within 1 h.

Earlier version: The intensity of the CD151 staining between cells is highly variable (Figure 5A and B). In the control, the mean CD151 intensity shows no trend over the 180 min time course, whereas with CytD the intensity diminishes after 180 min (Figure 5B). The decrease of the CD151 staining intensity points towards the possibility that, after CytD wash off, CD151 is more strongly internalized, presumably due to co-internalization with endocytosed PsVs. This idea is supported by the observation that, in particular at CytD/180 min, we occasionally observe CD151/PsV agglomerations (Figure 5A, see structures marked by arrows at CytD/180 min, for more examples see Supplementary Figure 5A). We did not study this issue systematically, but some of these structures have clear three-dimensional extension (see Supplementary Figure 5B for axial scans). Therefore, they are likely tubular structures filled with several PsVs, as previously described by electron microscopy (Schelhaas et al., 2012). We observe fewer of such structures at control/180 min, probably because cells have been actively interacting with PsVs for altogether 8 h, opposed to 3 h in the CytD treated cells. Hence, in the control, PsV endocytosis may be less synchronized compared to CytD explaining the CD151 diminishment at CytD/180 min (Figure 5B).

Revised version: The intensity of the CD151 staining between cells is highly variable (Figure 5A and B). In the control, the mean CD151 intensity shows no trend over the 180 min time course, whereas with CytD the intensity after 180 min is diminished in comparison to the control (Figure 5B). The decrease of the CD151 staining intensity points towards the possibility that, after CytD wash off, CD151 is more strongly internalized compared to the control, presumably due to increased co-internalization with endocytosed PsVs. This idea is supported by the observation that, in particular at CytD/180 min, we occasionally observe CD151/PsV agglomerations (Figure 5A, see structures marked by arrows at CytD/180 min, for more examples see Supplementary Figure 4A). We did not study this issue systematically, but some of these structures have clear three-dimensional extension (see Supplementary Figure 4B for axial scans). Therefore, they are likely tubular structures filled with several PsVs, as previously described by electron microscopy (Schelhaas et al., 2012). We observe fewer of such structures at control/180 min, probably because cells have been actively interacting with PsVs for altogether 8 h, opposed to 3 h in the CytD treated cells. Hence, after CytD wash off, PsV endocytosis may be more synchronized compared to control, explaining the CD151 staining intensity diminishment at CytD/180 min (Figure 5B).

Earlier version: For studying the association between PsVs (L1) and CD151, the PCC between the channels is calculated. The PCCs are around 0.1, with the exception of the CytD/0 min value that is significantly lower and even negative (Figure 5C, for PCC values of flipped images see Supplementary Figure 6). This reflects the somewhat mutual exclusion of the two stainings as already discussed above; PsV accumulations are at the cell periphery and the CD151 staining is mainly at the cell body and in filopodia. At 30 min after CytD, the PCC has reached the level of 0.1 as in the other conditions, suggesting that actin processes are required early for the PsV-CD151 association.

Revised version: For studying the association between PsVs (L1) and CD151, the PCC between the channels is calculated. The PCCs are around 0.1, with the exception of the 0 min/CytD value that is significantly lower and even negative (Figure 5C, for PCC values of flipped images see Supplementary Figure 5B and C). This reflects the in part mutual exclusion of the two stainings as already discussed above; PsV accumulations tend to be at the cell periphery and the CD151 staining is at the cell body and in filopodia. In any case, the lower PCC at 0 min/CytD suggests that without active recruitment less PsVs reach CD151. At 30 min after CytD, the PCC has reached the level of 0.1 as in the control, which is in line with the idea of fast recruitment as observed in Figure 4. To follow how the basal cell membrane is populated with PsVs over time, as additional analysis we determined the PsVs per µm2 in ROIs placed in the cell body region. At 0 min, CytD reduces the PsV density to 19 - 33%, albeit the effect is not significant, and at 180 min/CytD the same PsV density as in the control is reached (Supplementary Figure 6A and B). Overall, under CytD there was a trend towards less PsVs present (Supplementary Figure 6A and B). Hence, both Figure 5C and Supplementary Figure 6A and B suggest that active virion transport is required to reach efficiently the basal membrane.

Earlier version: One would expect that the translocation of PsVs to the cell surface increases the number of PsVs that are closely associated with CD151. As criteria for close association, we define a distance of ≤ 80 nm between PsV and CD151 maxima, a value which is close to the resolution limit of the used microscope (Finke et al., 2020). In the control, the fraction of PsVs closely associated with CD151, after correction for random background association (see Supplementary Figure 7), is 8.7% (Figure 5D, control). At 0 min after CytD, we start with a fraction of 5.5%, a value that increases to 12.7% in the next 60 min (CytD/60 min), although this change is not significant (Figure 5D, CytD).

Revised version: We next studied the fraction of PsVs that are closely associated with CD151. As criteria for close association, we define a distance of ≤ 80 nm between PsV and CD151 maxima because this value is close to the resolution limit of the used microscope (Finke et al., 2020). In the control, the fraction of PsVs closely associated with CD151 is around 10% (Figure 5D, control), after correction for random background association, for which we used a calibration line based on the same density of PsVs in flipped images (see Supplementary Figure 7). At 0 min after CytD, we start with a fraction of 5.5% of closely associated PsVs, that increases to 12.7% in the next 60 min, although this increase is not significant (Figure 5D, CytD). In any case, it is remarkable that in the absence of active recruitment (0 min/CytD) we find half the fraction of closely associates PsVs compared to the control. This indicates that PsVs and CD151 associate very early in the infection cascade, essentially in the moment the PsVs reach the edge of the cell body.

Earlier version: Altogether, we conclude that after a 5 h pre-incubation with PsVs (Figure 5D, control), about 10% of the PsVs are associated closely with CD151 in a time window of 180 min. CytD diminishes only the early association at 0 min (5.5%, Figure 5D, CytD). The data suggests that PsVs establish contact to CD151 assemblies early in the infection cascade (see also the increase in the PCC after CytD between 0 and 30 min, Figure 5C).

Revised version: In summary, we conclude that within 180 min after a 5 h pre-incubation with PsVs (Figure 5D, control), we encounter a type of steady-state situation (new PsVs are recruited whereas older ones disappear by endocytosis), in which about 10% of the PsVs are associated closely with CD151. CytD diminishes only the PsV-CD151 association at 0 min (5.5%, Figure 5D, CytD), suggesting that PsVs establish contact to CD151 early (see also the increase in the PCC between 0 min/CytD and 30 min/CytD, Figure 5C).

PsV association with HS

Next, we studied the association between PsVs and HS. As a reference staining for the cell body, we used Itgα6 that is not visible at cell protrusions. PsVs are visualized by click-chemistry and imaged in the confocal channel. HS and Itgα6 were stained by antibodies, and imaged at STED microscopic resolution. The three stainings were simultaneously recorded.

As shown in Figure 6A (green LUT), the Itgα6 staining results in a pattern of dense spots. The Itgα6 intensity does not change over time (Supplementary Figure 8E). The pattern of the HS staining (cyan LUT) and the overlap of HS with PsVs and Itgα6 are highly variable (Figure 6A).

Earlier version: Association between PsVs and HS.

Revised version: Association between PsVs and HS after CytD treatment.

(A) HaCaT cells were incubated with PsVs at 37 °C for 5 h, in the absence (Control, upper panels) or presence of 10 µg/ml CytD (CytD, lower panels). Afterwards, cells were washed and incubated without PsVs/CytD further for up to 180 min, before they were fixed and stained. PsVs (magenta LUT) were visualized by click-chemistry (6-FAM Azide) and indirect immunolabeling was used for HS (AlexaFluor 594; cyan LUT) and for Itgα6 (STAR RED; green LUT). The bottom rows show magnified views of the white boxes in the merged images. PsV-DNA staining was imaged in the confocal and HS and Itgα6 staining in the STED mode of a STED microscope. For analysis, we placed large ROIs onto the images that covered mainly the cell body but included parts of the cell periphery as well (see example in Supplementary Figure 8A). For (D), smaller ROIs covering only the cell body region were used. (B) The mean HS intensity plotted over time. (C) The PCC between PsV-DNA (magenta LUT) and HS (cyan LUT) over time. (D) The PCC between PsV-DNA (magenta LUT) and HS (cyan LUT) in the region of the cell body over time. (E) The fraction of PsVs (in percent) closely associating with HS (distance ≤ 80 nm) plotted over time (for background correction see Supplementary Figure 9). Two examples of PsVs (each marked by an asterisk) from the CytD/0 min condition are shown. The value in the upper left corner states the shortest distance (in nm) between the marked PsV and its next nearest HS maximum (marked by an arrow). Values are given as means ± SD (n = 3 biological replicates). Using the two-tailed, unpaired student’s t-test (n = 3 biological replicates), we analyzed in (B), (C) and (E) the statistical differences between the same time points of Control and CytD, and in (D) the difference between CytD/30 min and CytD/0 min or CytD/180 min (for details see methods).

Earlier version: CytD increases the intensity of HS (Figure 6B; also apparent when comparing the HS brightness in the upper and lower panels of Figure 6A; for an overview of the CytD/0min images see Supplementary Figure 4C). This increase of intensity is particularly notable at the 0 min time point, where the samples treated with CytD have a more than two-fold higher intensity and differ significantly from the control. Hence, in this experiment we reproduce the PsV/CytD mediated HS intensity increase observed in Figure 3.

Revised version: CytD increases the intensity of HS (Figure 6B; also apparent when comparing the HS brightness in the upper and lower panels of Figure 6A; for an overview of the CytD/0min images see Supplementary Figure 3C). This increase of intensity is particularly notable at the 0 min time point, where the samples treated with CytD have a more than two-fold higher intensity and differ significantly from the control. Hence, in this experiment we reproduce the PsV/CytD mediated HS intensity increase observed in Figure 3B.

Earlier version: In the control, the PCC between PsVs and HS is little above zero. The largest PCC is found at CytD/0 min, which reflects the finding that at this time point both PsVs and HS preferentially locate at the cell periphery (Figure 6C, for PCC values of flipped images see Supplementary Figure 8). Over time, the accumulated PsVs diminish due to translocation, which is accompanied by a PCC approaching zero, the same value as in the control (Figure 6C). Additionally, we analyze the PCC between PsVs and HS specifically in the cell body region, excluding the cell periphery (Figure 6D). For the control, the PCC in the cell body region ranges between 0 and 0.05. After CytD, we observe an increase in the PCC from 0 min to 30 min and a decrease from 30 min to 180 min (for PCC values of flipped images see Supplementary Figure 8).

Revised version: In the control, the PCC between PsVs and HS are low with no clear trend. The largest PCC is found at CytD/0 min, which reflects the finding that at this time point both PsVs and HS preferentially locate at the cell periphery (Figure 6C, for PCC values of flipped images see Supplementary Figure 8B and C). Over time, the accumulated PsVs diminish due to recruitment, which is accompanied by a PCC approaching zero, as in the control (Figure 6C, 180 min). Additionally, we analyze the PCC between PsVs and HS specifically in the cell body region, excluding the cell periphery (Figure 6D). For the control, the PCC in the cell body region ranges between 0 and 0.05. After CytD, we observe an increase in the PCC from 0 min to 30 min and a decrease from 30 min to 180 min (for PCC values of flipped images see Supplementary Figure 8F and G). When analyzing the PsVs closely associating with HS, at 0 min, CytD increases the fraction of PsVs associated with HS more than 3-fold. Over the next 180 min, this fraction gradually decreases until it is equal to the control value (Figure 6E, 180 min).

Removed in the revised version: When analyzing the PsVs closely associating with HS, at 0 min, CytD increases the fraction of PsVs associated with HS more than 3-fold. Over the next 180 min, this fraction gradually decreases until it is equal to the control value (Figure 6E).

Earlier version: Altogether, the analysis of the PCC between HS and PsVs (Figure 6C) along with the fraction of closely HS associated PsVs (Figure 6E) indicates that CytD treatment increases HS/PsV association at 0 min, likely because PsVs remain accumulated in the ECM. After CytD removal, this association diminishes. A transient increase in PCC is observed at 30 min specifically in the region of the cell body (Figure 6D) which likely reflects the translocation of HS-coated virions to the cell body. However, this association also diminishes over the subsequent 150 min, suggesting that an increasing number of PsVs are no longer associated with HS – likely due to progressive loss of the HS coat over time. Alternatively, PsVs just may disappear after internalization.

Revised version: Altogether, the analysis of the PCC between HS and PsVs (Figure 6C) along with the fraction of closely HS associated PsVs (Figure 6E) indicates that CytD treatment increases HS/PsV association at 0 min, likely because PsVs remain accumulated in the ECM. After CytD removal, this association diminishes. An increase in PCC is observed at 30 min specifically in the region of the cell body (Figure 6D) which likely reflects the recruitment of HS-coated virions to the cell body. However, as this association diminishes over the subsequent 150 min, an increasing fraction of PsVs is no longer associated with HS due to progressive loss of the HS coat over time, or alternatively, PsVs just may disappear after internalization.

Earlier version: Next, for each PsV we plot the ‘distance to the next nearest Itgα6 maximum’ against ‘the distance to the next nearest HS maximum’ (Figure 7). In the control, the distances remain unchanged over the entire observation time; 59% - 65% of the PsVs are at a short distance (< 250 nm) to both Itgα6 and HS (Figure 7B, black; see also Supplementary Table 1), 12 - 15% are at a short distance (< 250 nm) to HS but not to Itgα6 (> 250 nm) (Figure 7B, green; see also dashed green box in Figure 9 marking PsVs in the ECM). Also, the distance plots are in line with the idea that the distance patterns of PsVs in untreated cells do not change over the 180 min observation time (Figure 7C, please compare the upper distance plots from left to right). In cells treated with CytD, a larger fraction of PsVs is accumulated at the periphery at the 0 min time point, which is reflected in a larger fraction (46% as opposed to 12% in the control; Figure 7B, green; Supplementary Table 1) of PsVs with a short distance (< 250 nm) to HS but not to Itgα6. Over time, the PsVs in the CytD treated cells acquire a shorter distance to Itgα6 and a larger distance to HS (Figure 7C, lower row). The population with short distances to HS and large ones to Itgα6 strongly diminishes from 46% to 11% (Figure 7B, CytD, green). After 180 min, the distances are similar to the untreated control. Between CytD/60 min and CytD/180 min, the fraction of PsVs with a large distance (> 250 nm) to HS and a short distance (< 250 nm) to Itgα6, representing PsVs at the cell body without HS (see also dashed magenta box in Figure 9), increases from 9.6% to 25.2% (Figure 7B, magenta). This is inconsistent with the idea of endocytosis of HS-coated PsVs that would only diminish short distances to Itgα6 but not create long distances to HS, which is what we observe. The observation supports the idea that PsVs lose their HS coat after translocating to the cell surface, which is also supported by the transient increase in the PCC between HS and PsVs of the cell body at 30 min (Figure 6D).

Revised version: Next, for each PsV we plot the ‘distance to the next nearest Itgα6 maximum’ against ‘the distance to the next nearest HS maximum’ (Figure 7). In the control, the distances remain essentially unchanged over the entire observation time; 59% - 65% of the PsVs are at a short distance (< 250 nm) to both Itgα6 and HS (Figure 7B, black; see also Supplementary Table 1), 12 - 15% are at a short distance (< 250 nm) to HS but not to Itgα6 (> 250 nm) (Figure 7B, green; see also dashed green box in Figure 9 marking PsVs in the ECM). Regarding the distance plots of untreated cells in Figure 7 (please compare the upper distance plots from left to right), they suggest that the distance patterns do not change over the 180 min observation time. In cells treated with CytD, a larger fraction of PsVs is accumulated at the periphery at the 0 min time point, which is reflected in a larger fraction (46% as opposed to 12% in the control; Figure 7B, green; Supplementary Table 1) of PsVs with a short distance (< 250 nm) to HS but not to Itgα6. In the distance plots, it can be noticed that the PsVs in the CytD treated cells over time acquire shorter distances to Itgα6 and larger distances to HS (Figure 7C, lower row). The population with short distances to HS and large ones to Itgα6 strongly diminishes from 46% to 11% (Figure 7B, CytD, green). After 180 min, the distances are similar to the untreated control. Between CytD/60 min and CytD/180 min, the fraction of PsVs with a large distance (> 250 nm) to HS and a short distance (< 250 nm) to Itgα6, representing PsVs at the cell body without HS (see also dashed magenta box in Figure 9), increases from 9.6% to 25.2% (Figure 7B, magenta). This is inconsistent with the idea of endocytosis of HS-coated PsVs that would only diminish short distances to Itgα6 but not create long distances to HS, which is what we observe. The observation supports the idea that PsVs lose their HS coat after translocating to the cell surface, which is in line with the transient increase in the cell body PCC between HS and PsVs at 30 min (Figure 6D).

PsV-Itgα6 and PsV-HS distances over time.

Earlier version: (A) Definition of four PsV populations based on the PsV distances to Itgα6 and HS (the plot is taken from Control, 0 min and shown again in C). Dashed green rectangle, PsVs with a distance to HS < 250 nm and to Itgα6 > 250 nm. Dashed magenta rectangle, PsVs with a distance to HS > 250 nm and to Itgα6 < 250 nm. Dashed gray square, PsVs with a distance to HS > 250 nm and to Itgα6 > 250 nm. PsVs not included in the previous categories have a distance to HS < 250 nm and to Itgα6 < 250 nm (dashed black square). (B) From the PsVs analyzed in Figure 6, for the Control (left) and CytD (right) the PsV fraction size (in percent) of each population is illustrated. Shown are the means of three biological replicates. For means ± SD and statistical analysis see Supplementary Table 1. (C) For the Control (top) and CytD (bottom), for each PsV, we plotted the shortest distance to Itgα6 against the shortest distance to HS (pooling the three biological replicates; 3,043 – 4,080 PsVs per plot).

Revised version: (A) Definition of four PsV populations based on the PsV distances to Itgα6 and HS (please note that the data does not include distances of exactly 250 nm wherefore symbols as ≥ and ≤ are omitted; the plot is taken from Control, 0 min and shown again in C). Dashed green rectangle, PsVs with a distance to HS < 250 nm and to Itgα6 > 250 nm. Dashed magenta rectangle, PsVs with a distance to HS > 250 nm and to Itgα6 < 250 nm. Dashed gray square, PsVs with a distance to HS > 250 nm and to Itgα6 > 250 nm. PsVs not included in the previous categories have a distance to HS < 250 nm and to Itgα6 < 250 nm (dashed black square). (B) From the PsVs analyzed in Figure 6, for the Control (left) and CytD (right) the PsV fraction size (in percent) of each population is illustrated. Shown are the means of three biological replicates. For means ± SD and statistical analysis see Supplementary Table 1. (C) For the Control (top) and CytD (bottom), for each PsV, we plotted the shortest distance to Itgα6 against the shortest distance to HS (pooling the three biological replicates; 3,043 – 4,080 PsVs per plot).

Earlier version: The movement of HS towards the cell body after removal of CytD, which indirectly demonstrates that PsVs are coated with HS, is suggested by a shortening of the HS-Itgα6 distance over time (Supplementary Figure 8D). Moreover, the PCC between PsVs and HS in the cell body area increases (Figure 6D). Altogether, the data suggest that HS-coated PsVs migrate towards the cell body. A large fraction of translocating PsVs sheds HS within one hour after removing CytD (Figure 7C, lower panels, showing increasing PsV-HS distances). This is the time window where we notice the occurrence of endocytic PsV/CD151 structures (see above and Supplementary Figure 5).

Revised version: The recruitment of HS towards the cell body after removal of CytD, which indirectly demonstrates that PsVs are coated with HS, is suggested by a shortening of the HS-Itgα6 distance over 180 min only after CytD wash off (Supplementary Figure 8D). Together with the observation that the PCC between PsVs and HS in the cell body area increases (Figure 6D), the data suggest that HS-coated PsVs are recruited to the cell body. A fraction of PsVs sheds HS within one hour after removing CytD (Figure 7C, lower panels, showing increasing PsV-HS distances after 60 min). This is the time window in which we notice the occurrence of endocytic PsV/CD151 structures (see above and Supplementary Figure 4B).

Earlier version: Actin retrograde transport, which underlies the here observed virion transport, is the integrative result of three components (Schelhaas et al., 2008; Smith et al., 2008). On the one side, actin filaments in the cell periphery are pushed into the cell body area when actin polymerizes at their tips. The retrograde filament movement causes tension and is facilitated by F-actin degradation opposed to the side of F-actin growth. Moreover, the motor protein myosin II exerts force on actin filaments, pulling those towards the cell body. As CytD broadly interferes with F-actin dependent processes, we investigated the effects upon inhibition of only one of the three components, namely the myosin II mediated retrograde movement towards the cell body. Instead of CytD, we employed in the 5 h preincubation the myosin II inhibitor blebbistatin. For the control (0 min), we show in Figure 8A one example of a cell with comparatively many PsVs at the periphery (as mentioned above, the PsV pattern is highly variable) to better illustrate the difference to the PsV pattern occasionally seen with blebbistatin. After blebbistatin treatment (0 min), PsVs are still distal to the cell body but less dispersed than after CytD treatment, seemingly as if translocation started but stopped in the midst of the pathway (Figure 8A, blebbistatin). The PCC between PsVs and HS, like after CytD (Figure 6C), is elevated after blebbistatin, albeit the effect is not significant (Figure 8C). The cell body PCC, is not at 30 min (CytD) but already at 0 min elevated (compare Figure 6D to Figure 8D), which can be explained by partial translocation. This is further supported by the fact that only 8% of PsVs are closely associated with HS (Figure 8E; blebbistatin, 0 min) compared to 15% after CytD treatment (Figure 6E; 0 min).

Revised version: Actin retrograde transport underlies the filopodial virion transport and is the integrative result of three components (Smith et al., 2008; Schelhaas et al., 2008). On the one side, actin filaments in the cell periphery are pushed into the cell body area when actin polymerizes at their tips. The retrograde filament movement causes tension and is facilitated by F-actin degradation opposed to the side of F-actin growth. Moreover, the motor protein myosin II exerts force on actin filaments, pulling those towards the cell body. As CytD broadly interferes with F-actin dependent processes like filopodial transport and cell migration, we investigated the effects upon inhibition of only one component of actin retrograde flow, namely the myosin II mediated retrograde movement towards the cell body. Instead of CytD, we employed in the 5 h preincubation the myosin II inhibitor blebbistatin (Schelhaas et al., 2008). For the control (0 min), we show in Figure 8A one example of a cell with comparatively many PsVs at the periphery (as mentioned above, the PsV pattern is highly variable and also in the control we occasionally observe PsV accumulations) to better illustrate the difference to the PsV pattern occasionally seen with blebbistatin. After blebbistatin treatment (0 min), PsVs do not reach the central cell body as in the control, but are less dispersed than after CytD treatment, seemingly as if recruitment started but stopped in the midst of the pathway (Figure 8A, blebbistatin). The PCC between PsVs and HS, like after CytD (Figure 6C), is elevated after blebbistatin, albeit the effect is not significant (Figure 8C). The cell body PCC is not, as under CytD, at 30 min (Figure 6D) but already at 0 min elevated (compare Figure 6D to Figure 8D), which can be explained by partial recruitment. This is further supported by the fact that only 8% of PsVs are closely associated with HS (Figure 8E; blebbistatin, 0 min) compared to 15% after CytD treatment (Figure 6E; 0 min). Furthermore, after 0 min PsV incubation with blebbistatin we observe no effect on the HS intensity (compare Figure 8B to Figure 3B and Figure 6B). Hence, in contrast to CytD, blebbistatin does not preserve the PsVs in the ECM where they associate with HS, but ongoing actin polymerization may push actin filaments along with PsVs towards the cell body.

Association between PsVs and HS after blebbistatin treatment.

Earlier version: (A) HaCaT cells were incubated with PsVs at 37 °C for 5 h, in the absence (Control, upper panels) or presence of 30 µM blebbistatin (Blebbistatin, lower panels). Afterwards, cells were washed and incubated without PsVs/blebbistatin further for up to 180 min, before they were fixed and stained. Immunofluorescence was used for L1 (STAR GREEN; magenta LUT) and HS (AlexaFluor 594; cyan LUT) staining. F-actin was stained by iFluor647-labelled phalloidin (green LUT). The bottom rows show magnified views of the white boxes in the merged images. PsVs and F-actin staining were imaged in the confocal and HS staining in the STED mode of a STED microscope. For analysis, we placed large ROIs onto the images that covered mainly the cell body but included parts of the cell periphery as well. For (D), smaller ROIs covering only the cell body region were used. (B) Mean HS intensity over time. (C) PCC between PsV-L1 (magenta LUT) and HS (cyan LUT) over time. (D) PCC between PsV-L1 (magenta LUT) and HS (cyan LUT) in the region of the cell body over time. (E) The fraction of PsVs (in percent) closely associating with HS (distance ≤ 80 nm) plotted over time (for background correction see Supplementary Figure 11). Two examples of PsVs (each marked by an asterisk) from the blebbistatin/0 min condition are shown. The value in the upper left states the shortest distance between the PsV and its next nearest HS maximum (marked by an arrow) in nm. Values are given as means ± SD (n = 3 biological replicates). Statistical differences between the same time points of Control and Blebbistatin were analyzed by using the two-tailed, unpaired student’s t-test (n = 3), but the analysis yielded no p-values below 0.05.

Revised version: (A) HaCaT cells were incubated with PsVs at 37 °C for 5 h, in the absence (Control, upper panels) or presence of 30 µM blebbistatin (Blebbistatin, lower panels). Afterwards, cells were washed and incubated without PsVs/blebbistatin further for up to 180 min, before they were fixed and stained. Immunofluorescence was used for L1 (STAR GREEN; magenta LUT) and HS (AlexaFluor 594; cyan LUT) staining. F-actin was stained by iFluor647-labelled phalloidin (green LUT). The bottom rows show magnified views of the white boxes in the merged images. PsVs and F-actin staining were imaged in the confocal and HS staining in the STED mode of a STED microscope. For analysis, we placed large ROIs onto the images that covered mainly the cell body but included parts of the cell periphery as well. For (D), smaller ROIs covering only the cell body region were used. (B) Mean HS intensity over time. (C) PCC between PsV-L1 (magenta LUT) and HS (cyan LUT) over time (for control with flipped images see Supplementary Figure 10A and B). (D) PCC between PsV-L1 (magenta LUT) and HS (cyan LUT) in the region of the cell body over time (for control with flipped images see Supplementary Figure 10C and D). (E) The fraction of PsVs (in percent) closely associating with HS (distance ≤ 80 nm) plotted over time (for background correction see Supplementary Figure 11). Two examples of PsVs (each marked by an asterisk) from the blebbistatin/0 min condition are shown. The value in the upper left states the shortest distance between the PsV and its next nearest HS maximum (marked by an arrow) in nm. Values are given as means ± SD (n = 3 biological replicates). Statistical differences between the same time points of Control and Blebbistatin were analyzed by using the two-tailed, unpaired student’s t-test (n = 3), but the analysis yielded no p-values below 0.05.

Added in the revised version: PsV binding in the absence of a diffusion barrier

Added in the revised version: Throughout all experiments, we observe at 0 min/CytD only few PsVs at the basal membrane (Figure 1A, Supplementary Figure 6A and B; see also PCC at 0 min between PsVs an CD151 in Figure 5C), suggesting that in the absence of active recruitment the access to the basal membrane via passive diffusion is limited. We wondered, how many PsVs may bind to the cell membrane without a diffusion barrier? For this reason, we incubated EDTA detached HaCaT cells in suspension with PsVs for 1 h at 4 °C, followed by re-attachment for 1 h. Under these conditions, we find, despite of a shorter incubation time (1 h versus 5 h), a roughly 3-fold larger PsV density (1.7 PsVs/µm2 (Supplementary Figure 6D)) than the highest density observed in the other experiments. However, it should be noted that values of the different experiments cannot be directly compared. Aside from the different treatments, another difference lies in the size of the imaged membrane. The re-attachment of cells is not complete after 1 h (compare size of adhered membranes in Supplementary Figure 6A and 1A), wherefore the membranes are likely strongly ruffled, which results in the underestimation of the membrane area. As a result, we overestimate the PsVs per µm2 adhered membrane (please note that we cannot re-attach cells for longer times as we then lose PsVs due to endocytosis). In any case, the experiment suggests that PsVs bind more efficiently to membrane surface receptors without a diffusion barrier. We conclude that in our assay PsVs cannot readily bypass the active PsV recruitment by diffusing directly to the basal cell membrane, which is plausible, because to make this happen a 55 nm large PsV must diffuse through the narrow gap between glass-coverslip and adhered cell.

Removed in the revised version: Furthermore, after 0 min PsV incubation with blebbistatin we observe no effect on the HS intensity (compare Figure 8B to Figure 3B and Figure 6B). Hence, in contrast to CytD, blebbistatin does not trap the PsVs in the ECM where they associate with HS, but ongoing actin polymerization pushes actin filaments along with PsVs towards the cell body.

Discussion

Earlier version: In this study, we investigate the early events of HPV16 infection preceding virus endocytosis. Our findings reveal that actin dynamics of filopodia, rather than passive virion diffusion, are crucial for releasing HS-coated virus particles from the adhesive ECM and translocating them to the cell body. When actin-dependent processes are inhibited by CytD, HS-coated virus particles remain trapped in the ECM, allowing for a more synchronized observation of the infection process. Once the transfer is no longer blocked, PsVs rapidly migrate toward the cell body, associate with CD151, and shed their HS-coat before endocytosis.

Revised version: In this study, we investigate the early events in HPV16 infection occurring at the cell surface. Our findings reveal that in our assay actin dependent mechanisms, rather than passive virion diffusion, are crucial for recruiting HS-coated virus particles from the adhesive ECM to the basal cell membrane. PsVs associate early with CD151, and shed their HS-coat before endocytosis. Moreover, we propose that reversible blocking of active recruitment by CytD allows for a more synchronized observation of the steps on the cell surface.

Earlier version: A reversible blocking of PsV translocation from the ECM to the cell body by CytD

Revised version: A reversible CytD blocking of PsV recruitment from the ECM to the basal cell membrane

Earlier version: Many viruses use cell surface HS for primary attachment, including herpes simplex virus type 1, human cytomegalovirus, human immunodeficiency virus type 1, adenovirus type 2, dengue virus, hepatitis B virus, and vaccinia virus (Cagno et al., 2019; Giroglou et al., 2001; Tian et al., 2021). HPVs could be different as they bind as well to HS of the extracellular basement membrane, at least in vivo during the wounding and healing processes required for infection (Day and Schelhaas, 2014; Kines et al., 2009; Ozbun and Campos, 2021; Schiller et al., 2010).

Revised version: Many viruses use cell surface HS for primary attachment, including herpes simplex virus type 1, human cytomegalovirus, human immunodeficiency virus type 1, adenovirus type 2, dengue virus, hepatitis B virus, and vaccinia virus (Cagno et al., 2019; Giroglou et al., 2001; Tian et al., 2021).

Added in the revised version: HPVs could be different as they bind as well to HS of the extracellular basement membrane, at least in vivo during the wounding and healing processes required for infection (Day and Schelhaas, 2014; Kines et al., 2009; Ozbun and Campos, 2021; Schiller et al., 2010).

Earlier version: By treating cells with CytD, we block actin-mediated transport, leading to a 6-fold accumulation of PsVs at the cell periphery after 5 hours (Figure 4B). After CytD removal, infectivity is only moderately reduced (29%) (Figure 1E), suggesting that the treatment is reversible and only slightly harmful, if at all.

Revised version: By treating cells with CytD, we block actin-mediated recruitment, leading to a 6-fold accumulation of PsVs at the cell periphery after 5 hours (Figure 4B). After CytD removal, infectivity is reduced by 29% (Figure 1E). Some reduction is to be expected, as PsVs have ≈20% less time to complete infection as compared to cells that were not treated with CytD. Hence, the infection assay suggests that the treatment is largely reversible and only slightly harmful, if at all. However, the luciferase infection assay does not distinguish between actively recruited PsVs and PsVs that bind passively by diffusion to the upper membrane. The latter fraction likely dominates the total infection rate and should be less affected by CytD than the fraction of actively recruited PsVs. Therefore, if the infection pathway of a small fraction of actively recruited PsVs is irreversibly inhibited, we may not be able to detect this effect on the background of unaffected passively binding PsV.

Added in the revised version: Upon CytD washout, PsVs approach the basal membrane with a half-time of ≈15 minutes, as suggested by the diminishment of accumulated PsVs at the periphery (Figure 4B). The half-time could be longer if cell spreading is also underlying the translocation of PsVs onto the cell body. However, we assume that this is rather unlikely, as cell spreading would increase the PCC between PsVs and F-actin under a condition where PsVs are not-primed (and therefore not actively recruited) but cell spreading occurs, which is not the case in Figure 2B and D (CytD/leupeptin).

Added in the revised version: In Figure 4, we observe fast diminishment of the peripheral PsVs within 15 min. It is in principle possible that after CytD removal ECM-accumulated PsVs are merely washed off. However, after only 15 min, the effect of this could not have been very strong, considering the still 6-fold accumulation of PsVs after the 5 h incubation with CytD.

Earlier version: Upon CytD washout, PsVs translocate to the cell body with a half-time of ≈15 minutes, as suggested by the diminishment of accumulated PsVs at the periphery. As mentioned above, the half-time could be longer if cell spreading is in part responsible for the translocation of PsVs onto the cell body. However, we assume that this is rather unlikely, as cell spreading would increase the PCC between PsVs and F-actin under a condition where filopodia mediated transport is blocked but not cell spreading, which is not the case (Figure 2B and D, CytD/leupeptin). Within this 15 min time frame, PsVs associate more strongly with CD151. It is in principle possible that along with CytD removal ECM-accumulated PsVs are merely washed off. However, after only 15 min, the effect of this could not have been very strong, considering the still 6-fold accumulation of PsVs after the 5 h incubation with CytD. Furthermore, CytD pre-treatment appears to synchronize uptake: We observe endocytic structures after CytD treatment more frequently (for examples see Supplementary Figure 5) and the CD151 intensity diminishes over time only after the CytD treatment (Figure 5B). We therefore conclude that the processing of the PsVs in the ECM is contributing to desynchronization and is largely completed after the 5 h pre-incubation with CytD. We estimate that about 6-fold more PsVs (Figure 2B) translocate to the cell surface in a coordinated fashion after CytD removal, which likely enhances the number of endocytic events compared to untreated controls.

Revised version: Furthermore, CytD pre-treatment appears to synchronize uptake: We observe endocytic structures after CytD treatment more frequently and the CD151 intensity diminishes over time only after the CytD treatment (Figure 5B). We estimate that, compared to control, about 6-fold more PsVs (Figure 4B) approach the cell surface in a coordinated fashion after CytD removal, which likely enhances the number of endocytic events compared to untreated cells. We therefore propose that the processing of the PsVs in the ECM is contributing to desynchronization and is largely completed after the 5 h pre-incubation with CytD.

The role of actin-driven virion transport

Added in the revised version: The strong electrostatic interaction between HPV and HS implies that some processing — via HS cleavage or structural capsid changes — is a prerequisite for the release of virions from the ECM.

Earlier version: The strong electrostatic interaction between HPV and HS implies that some processing — via HS cleavage or structural capsid changes — is a prerequisite for the release of virions from the ECM. CytD likely does not inhibit HS processing (Jamieson et al., 2021; see also Supplementary Figure 2) nor is it likely to inhibit KLK8. Therefore, we assume that PsVs are primed after 5 h of incubation with CytD. They may remain associated via weaker interactions, and actin-driven transport is required for pulling the HS-coated virions out of the sticky matrix towards the cell body. Sticky PsV aggregations that are released from HeLa cell surface after binding and blocking of the subsequent entry pathway have been observed earlier and might reflect similar HS/PsV-stuctures (Mikuličić et al., 2025). Alternatively, an ECM bound virion is not exposed to all processing factors at the same time. In this case, migration through the ECM is additionally required to establish, one by one, contact with all processing enzymes involved. Once the fully processed virions have reached the cell surface, detachment from the ECM may be facilitated by new bonds established between the PsVs and cell surface receptors. Now, the PsVs and cell surface receptors move towards the central cell body by intracellular dynamics which results in CD151 accumulation at virus binding sites and ultimately endocytosis (see Figure 5 as well as Supplementary Figure 5).

Revised version: CytD likely does not inhibit HS processing (Jamieson et al., 2021) nor is it likely to inhibit KLK8. Therefore, we assume that, although PsVs are primed after 5 h of incubation with CytD, they remain in the ECM. They may be associated via weaker interactions, and actin-driven transport is required for pulling the HS-coated virions out of the sticky matrix towards the cell body. Sticky PsV aggregations that are released from HeLa cell surface after binding and blocking of the subsequent entry pathway have been observed earlier and might reflect similar HS/PsV-stuctures (Mikuličić et al., 2025). It is possible that an ECM bound virion is not exposed to all required processing factors at the same time. In this case, migration through the ECM enables for contact with all processing enzymes. Once the fully processed virions have reached the cell surface, detachment from the ECM may be facilitated by new bonds established between the PsVs and cell surface receptors. Now, the PsVs and cell surface receptors move towards the center of the cell body by intracellular dynamics which results in CD151 accumulation at virus binding sites (see Figure 5 as well as Supplementary Figure 4) and ultimately endocytosis.

Earlier version: Both CytD and blebbistatin increase PsV-HS colocalization, but blebbistatin allows partial translocation (compare Figure 6C and Figure 8C as well as Figure 6D and Figure 8D), indicating that actin retrograde flow initiates movement, while myosin II contributes to sustained transport. As a result, virus movement can proceed to some extent after blebbistatin treatment. However, the subsequent blocking of transport at later time points reflects the inhibition of myosin II, which eventually disrupts the cytoskeletal dynamics and force generation that is necessary for sustained transport. In contrast, CytD directly interferes with actin polymerization, leading to an earlier and more persistent blocking of virus transport due to structural disruption of the filopodial tracks themselves.

Revised version: Both CytD and blebbistatin increase PsV-HS colocalization, but blebbistatin allows partial recruitment (compare Figure 6C and Figure 8C as well as Figure 6D and Figure 8D), indicating that actin retrograde flow initiates movement, while myosin II contributes to sustained transport. As a result, virus movement can proceed to some extent after blebbistatin treatment. However, the subsequent blocking of transport at later time points reflects the inhibition of myosin II, which eventually disrupts the cytoskeletal dynamics and force generation that is necessary for sustained transport. In contrast, CytD directly interferes with actin polymerization, leading to an earlier and more persistent blocking of virus transport due to structural disruption of the filopodial tracks themselves.

Earlier version: Altogether, our findings support earlier reports showing that HPV utilizes filopodia for cell entry, migrating at several micrometers per minute (Schelhaas et al., 2008; Smith et al., 2008). This translocation is fast compared to the entire infection process and therefore cannot largely contribute to the asynchronous HPV uptake. Since we use polarized cells, PsVs hardly bind to the cell surface directly (Supplementary Figure 1). Therefore, the actin-driven virion transport plays a decisive role in HPV infection of isolated HaCaT cells, which cannot be modelled in nonpolarized cells, where PsVs can bind to the entire cell surface.

Revised version: Altogether, our findings support earlier reports showing that HPV utilizes filopodia for cell entry, migrating at several micrometers per minute (Schelhaas et al., 2008; Smith et al., 2008). This translocation is fast compared to the entire infection process and therefore cannot largely contribute to the asynchronous HPV uptake.

Added in the revised version: The analyzed PsVs hardly bind to the basal cell surface directly by diffusion (Supplementary Figure 6, compare PsV maxima density at 0 min/CytD in A and B to C). Therefore, the actin-driven virion transport would play a decisive role in HPV infection if cells would form a monolayer with a disruption at which ECM is present and that is approached by PsVs, a scenario similar to in vivo infection. In addition, cell migration could establish contact between PsVs and the cell surface.

HaCaT cells as model system in HPV infection

Added in the revised version: Inhibition of HPV16 PsV transport diminishes infection only of subconfluent but not of confluent HaCaT cells by about 50% (Schelhaas et al., 2008). Therefore, actin-dependent translocation along protrusions may be dispensable for infection at a high cell density (Schelhaas et al., 2008) and merely increases the exposure of cells to virions (Smith et al., 2008) that can readily bind to the upper cell membrane. We are not aware of a PsV translocation mechanism from the upper to the basal membrane. Therefore, in our assay, PsVs bound to the upper membrane are not expected to show up at the basal membrane. Comparing 0 min of control and CytD (Supplementary Figure 6A and B), we find that compared to the control 19 - 33% of the PsVs reach the basal membrane in the absence of active transport, or in other words, most likely by passive diffusion. Actually, the range from 19 – 33% must be a strong overestimate as PsVs in the control are in transit and many actively recruited PsVs are already internalized during the 5 h incubation period. For this reason, we propose that most likely much less than 10% of the PsVs reach the basal membrane by diffusion. Moreover, in the absence of the diffusion barrier, the density of bound PsVs is strongly increased (Supplementary Figure 6D), showing indirectly that at the basal membrane the binding sites are difficult to access without active recruitment. Taken together, we propose the large majority of PsVs analyzed in our assay are ECM bound and actively recruited to the basal cell membrane.

Earlier version: Inhibition of HPV16 PsV transport diminishes infection only of subconfluent but not of confluent HaCaT cells by about 50% (Schelhaas et al., 2008). Therefore, actin-dependent translocation along protrusions may be dispensable for infection at a high cell density (Schelhaas et al., 2008) and merely increases the exposure of cells to virions (Smith et al., 2008). However, this increased exposure could be relevant in vivo, as wounding of the epidermis results in upregulation of filopodia formation (Vasioukhin et al., 2000). Filopodia usage not only facilitates infection but also increases the likelihood of virions to reach their target cells during wound healing, namely the filopodia-rich basal dividing cells. In fact, several types of viruses exploit filopodia during virus entry (Chang et al., 2016), hinting at the possibility that for HPV and other types of viruses actin-driven virion transport may play a more important role than it is currently assumed. If this is the case, sub-confluent HaCaT cells, or even better single HaCaT cells, would be an ideal model system for the study of these very early infection steps that involve ECM attachment and subsequent filopodia-dependent transport. As shown in Supplementary Figure 1, HaCaT cells have many binding sites for the HPV16 PsVs. However, as they are polarized and the binding receptors are only at the basal membrane, they remain relatively inaccessible by diffusion. Therefore, the ECM binding that is also observed in vivo (Day and Schelhaas, 2014) and subsequent transport via filopodia are used upon infection of HaCaT cells that locate at the periphery of cell patches. Here, PsVs bind to the ECM which strongly enhances infection of primary keratinocytes (Bienkowska-Haba et al., 2018). In contrast, HPV can readily bind to HSPGs on the cell surface of nonpolarized cells, and by this bypasses ECM mediated virus priming and the filopodia dependency. We propose that HaCaT cells are a valuable system for studying the very early events in HPV infection that allows for dissecting capsid interaction with ECM resident priming factors and cell surface receptors.

Revised version: The increased virion exposure via active recruitment could be relevant in vivo, as wounding of the epidermis results in upregulation of filopodia formation (Vasioukhin et al., 2000). Filopodia usage not only facilitates infection but also increases the likelihood of virions to reach their target cells during wound healing, namely the filopodia-rich basal dividing cells. In fact, several types of viruses exploit filopodia during virus entry (Chang et al., 2016), hinting at the possibility that for HPV and other types of viruses actin-driven virion transport may play a more important role than it is currently assumed. In this scenario, sub-confluent HaCaT cells, or even better single HaCaT cells, would be an ideal model system for the microscopic study of these very early infection steps that involve ECM attachment and subsequent active recruitment, as supposed to occur during in vivo infection of basal keratinocytes after binding of virions to the basement membrane (Bienkowska-Haba et al., 2018; Day and Schelhaas, 2014; Kines et al., 2009; Schiller et al., 2010). In contrast, in biochemical infection assays, virions diffusing to HSPGs on the cell surface, and by this bypassing active recruitment, are assayed together with the actively recruited virions. Should cells secrete little ECM and are grown to confluency, the passively binding virions are supposed to strongly dominate the infection rate in a biochemical infection assay.

Earlier version: Translocating PsVs are decorated with HS which they lose over time

Revised version: Recruited PsVs are decorated with HS which they lose over time

CytD treatment increases the HS signal when PsVs are present (Figure 3B and 6B). A relationship between PsVs and HS cleavage products has been previously suggested (Ozbun and Campos, 2021; Surviladze et al., 2012), and we propose that our data is in line with these previous observations. It should be noted that under conditions without the 5 h CytD blocking, PsVs may only briefly reside in the ECM and as a result are less decorated with HS. Hence, under our assay conditions the HS decoration of virions may be enhanced in comparison to other assays.

Earlier version: At the 0 min time point post-CytD, 15% of the PsVs are closely associated with HS (Figure 6E), but this fraction declines to 4.3% by 180 min. Considering the important role of HS as a primary attachment site, the fraction of 15% appears unexpectedly small. There are multiple possible explanations for underestimating the number of PsVs closely associated with HS. One contributing factor may be that accumulated PsVs are not well-resolved (see above), which results in an underestimation of HS associated PsVs in particular at CytD/0 min. Other explanations include that the HS antibody does not reach all HS epitopes in the dense matrix. Yet another explanation could be that PsV binding to HS and antibody binding to HS is sometimes mutually exclusive, a hypothesis which is supported by data showing that the HPV capsid possesses multiple HS binding sites (Dasgupta et al., 2011; Richards et al., 2013). Together, our data suggests that PsVs initially translocate with HS attached and then shed it, potentially either upon receptor binding or further capsid rearrangement. These observations are in line with a previous study showing reduced HS colocalization during productive infection (Selinka et al., 2007).

Revised version: At the 0 min time point post-CytD, 15% of the PsVs are closely associated with HS, but this fraction declines to 4.3% within 180 min (Figure 6E). Considering the important role of HS as a primary attachment site, the fraction of 15% appears unexpectedly small. There are multiple possible explanations for underestimating the number of PsVs closely associated with HS. One factor may be that accumulated PsVs are not well-resolved (see above), which results in an underestimation of HS associated PsVs in particular at CytD/0 min. Other explanations include that the HS antibody does not reach all HS epitopes in the dense matrix. Yet another explanation could be that PsV binding to HS and antibody binding to HS is sometimes mutually exclusive, a hypothesis which is supported by data showing that the HPV capsid possesses multiple HS binding sites (Dasgupta et al., 2011; Richards et al., 2013). Finally, not all PsVs bound to the ECM are expected to bind to HS but could also bind to laminin 332 (Culp et al., 2006). Together, our data suggests that PsVs initially are recruited with HS attached to them and then shed it, potentially either upon receptor binding or further capsid rearrangement. These observations are in line with a previous study showing reduced HS colocalization during productive infection (Selinka et al., 2007).

Earlier version: The initial HS-coat may facilitate the formation of a complex with PsVs and cell surface receptors such as integrins or growth factor receptors (GFRs) (Ballut et al., 2013; Pellegrini, 2001). This could facilitate receptor clustering and engagement as shown previously for coronavirus infection (Bugatti et al., 2025; Clausen et al., 2020; Zhang et al., 2023), as well as the binding of signaling molecules to the cell surface (Miladinova et al., 2022). Here, CD151 may support receptor clustering by binding to both integrins and GFRs (Hemler, 2005; Mikuličić et al., 2019). In addition, earlier studies suggested that cofactors like laminin-332 and growth factors further stabilize these interactions (Culp et al., 2006; Mikuličić et al., 2019; Richards et al., 2014; Surviladze et al., 2012). Furthermore, the HPV capsid has multiple binding sites for distinct functions (Richards et al., 2013) which supports the idea that HS fulfills several roles during viral entry — first by promoting capsid attachment, then inducing conformational changes, and finally by assisting in receptor engagement and clustering. However, HS cleavage alone is not sufficient for ECM release; active transport and capsid priming remain essential (see leupeptin effect in Figure 2B and D). Diffusion may account for some virion movement, but filopodia-mediated transport dominates in this context.

Revised version: The initial HS-coat may facilitate the formation of a complex with PsVs and cell surface receptors such as integrins or growth factor receptors (GFRs) (Ballut et al., 2013; Pellegrini, 2001). This could facilitate receptor clustering and engagement as shown previously for coronavirus infection (Bugatti et al., 2025; Clausen et al., 2020; Zhang et al., 2023), as well as the binding of signaling molecules to the cell surface (Miladinova et al., 2022). Here, CD151 may support receptor clustering by binding to both integrins and GFRs (Hemler, 2005; Mikuličić et al., 2019). In addition, earlier studies suggested that cofactors like laminin-332 and growth factors further stabilize these interactions (Culp et al., 2006; Mikuličić et al., 2019; Richards et al., 2014; Surviladze et al., 2012). Furthermore, the HPV capsid has multiple binding sites for distinct functions (Richards et al., 2013) which supports the idea that HS fulfills several roles during viral entry — first by promoting capsid attachment, then inducing conformational changes, and finally by assisting in receptor engagement and clustering. However, HS cleavage alone is not sufficient for ECM release; active transport and capsid priming remain essential (see leupeptin effect in Figure 2B and D).

Earlier version: PsV-CD151 association is an early contact at the cell surface

Revised version: PsV-CD151 association occurs early at the cell surface

Earlier version: CD151 is a known mediator of HPV entry (Scheffer et al., 2013; Spoden et al., 2008). Our data shows that about 10% of the PsVs are closely associated with CD151 at all post-translocation time points. This fraction is lower at 0 min post-CytD (Figure 5D, see also lower PCC in Figure 5C), indicating that actin disruption impairs initial CD151 engagement. Hence, the data suggest that PsVs establish contact to CD151 assemblies shortly after translocation to the cell surface, and from then on, the fraction of closely associated PsVs remains constant. At first glance this implies that, after the initial formation of the PsV-CD151 assemblies, they do not change in nature. However, we observe agglomerated CD151 maxima (locally patched maxima) close to PsVs at later time points that likely are endocytic structures (see arrow in Figure 5A, Supplementary Figure 5). This agglomeration process is not detected in our analysis as we only determine the next nearest CD151 assembly. Some further CD151 reorganization into larger entry platforms has been previously observed (Florin and Lang, 2018). This is in accordance with the observation that virions enter the cell in crowds with many of them occupying one endocytic organelle (Schelhaas et al., 2012).

Revised version: CD151 is a known mediator of HPV entry (Scheffer et al., 2013; Spoden et al., 2008). Our data shows that about 10% of the PsVs are closely associated with CD151 at all time points. This fraction is lower at 0 min post-CytD (Figure 5D, see also lower PCC in Figure 5C), indicating that actin disruption impairs initial CD151 engagement. Hence, the data suggest that PsVs establish contact to CD151 assemblies already during recruitment to the cell surface, and from then on, the fraction of closely associated PsVs remains constant. At first glance this implies that, after the initial formation of the PsV-CD151 assemblies, they do not change in nature. However, we observe agglomerated CD151 maxima (locally patched maxima) close to PsVs at later time points that likely are endocytic structures (see arrow in Figure 5A and Supplementary Figure 4). This agglomeration process is not detected in our analysis as we only determine the next nearest CD151 maximum. Some further CD151 reorganization into larger entry platforms has been previously observed (Florin and Lang, 2018). This is in accordance with the observation that virions enter the cell in crowds with many of them occupying one endocytic organelle (Schelhaas et al., 2012).

Earlier version: For the PsVs closely associated with CD151, one may have expected a larger fraction than 10%. Given that PsVs are not supposed to bind directly to CD151 and that CD151 is likely part of a larger membrane structure, like a tetraspanin enriched microdomain, our ≤ 80 nm distance criteria may be too short and by this we underestimate the association frequency. Altogether, we propose that an early contact between PsVs and the cell surface involves an assembly defined by CD151 (Figure 9).

Revised version: For the PsVs closely associated with CD151, one may have expected a larger fraction than 10%. Given that PsVs are not supposed to bind directly to CD151 and that CD151 is likely part of a larger membrane structure, like a tetraspanin enriched microdomain, our ≤ 80 nm distance criteria may be too short and by this we underestimate the associated fraction. Altogether, we propose that an early contact between PsVs and the cell surface involves CD151 (Figure 9).

Earlier version: Model of ECM events, actin transport, and receptor engagement in HPV16 entry.

Revised version: Model of ECM events, recruitment, and receptor engagement in HPV16 entry.

Earlier version: (i) During 5 h of incubation with CytD, the PsVs bind to HS of the ECM, are primed, and become coated with HS cleavage products, enabling them for cell surface receptor engagement. (ii) After CytD removal, within 15 min, HS-decorated viruses move along filopodia to the cell body and associate with CD151 assemblies (completed within 30 min), likely located at the filopodia. (iii) Eventually, they lose their HS coat, and individual HPV16-CD151 assemblies agglomerate into larger structures (platforms), which are subsequently endocytosed. Dashed rectangles mark PsVs representing populations as defined in Figure 7. Dashed green rectangle, PsVs with a distance to HS < 250 nm and to Itgα6 > 250 nm. Dashed black rectangle, PsVs with a distance to HS < 250 nm and to Itgα6 < 250 nm. Dashed magenta rectangle, PsVs with a distance to HS > 250 nm and to Itgα6 < 250 nm.

Revised version: (i) During 5 h of incubation with CytD, the PsVs bind to HS of the ECM, are primed, and become coated with HS cleavage products, enabling them for cell surface receptor engagement. After CytD removal, within 15 min, HS-decorated viruses are actively recruited to the cell body and (ii) associate with CD151 assemblies (completed within 30 min). (iii) Eventually, they lose their HS coat, and individual HPV16-CD151 assemblies agglomerate into larger structures (platforms), which are subsequently endocytosed. Dashed rectangles mark PsVs belonging to populations as defined in Figure 7. Dashed green rectangle, PsVs with a distance to HS < 250 nm and to Itgα6 > 250 nm. Dashed black rectangle, PsVs with a distance to HS < 250 nm and to Itgα6 < 250 nm. Dashed magenta rectangle, PsVs with a distance to HS > 250 nm and to Itgα6 < 250 nm.

Integration of our data into the HPV infection cascade

Earlier version: HPV infection is the result of several steps, starting with the initial binding of virions via electrostatic and polar interactions (Dasgupta et al., 2011) to the primary attachment site HS (Richards et al., 2013), which induces capsid modification (Cerqueira et al., 2015; Feng et al., 2024) and HS cleavage (Surviladze et al., 2015), enabling the virion to be released from the ECM or the glycocalyx. Next, virions bind to the cell surface to a secondary receptor complex that forms over time, and become internalized via endocytosis, before they are trafficked to the nucleus (Mikuličić et al., 2021; Ozbun and Campos, 2021). Regarding the transition from the primary attachment site to cell surface binding, as already outlined in the introduction, two models are discussed. In one model, proteases cleave the capsid proteins. After priming, the capsids are structurally modified and the virion can dissociate from its HS attachment site. It has been suggested that capsid priming is mediated by KLK8 (Cerqueira et al., 2015) and furin (Richards et al., 2006). In our system, KLK8 inhibition blocks PsV transport, while furin inhibition has some effect that, however, cannot be detected in this analysis (Figure 2) suggesting furin engagement at later steps in the infection cascade. This is in line with earlier in vitro studies on the role of cell surface furin (Day et al., 2008; Day and Schiller, 2009; Surviladze et al., 2015). In any case, our results align with both models of ECM detachment: one involving HS cleavage (HS co-transfer) and another involving capsid modification (by e.g., KLK8).

Revised version: HPV infection is the result of several steps, starting with the initial binding of virions via electrostatic and polar interactions (Dasgupta et al., 2011) to the primary attachment site HS (Richards et al., 2013), which induces capsid modification (Cerqueira et al., 2015; Feng et al., 2024) and HS cleavage (Surviladze et al., 2015), enabling the virion to be released from the ECM or the glycocalyx. Next, virions bind to the cell surface to a secondary receptor complex that forms over time, and become internalized via endocytosis, before they are trafficked to the nucleus (Mikuličić et al., 2021; Ozbun and Campos, 2021). Regarding the transition from the primary attachment site to cell surface receptor binding, as already outlined in the introduction, two models are discussed. In one model, proteases cleave the capsid proteins. After priming, the capsids are structurally modified and the virion can dissociate from its HS attachment site. It has been suggested that capsid priming is mediated by KLK8 (Cerqueira et al., 2015) and furin (Richards et al., 2006). In our system, KLK8 inhibition blocks PsV transport, while furin inhibition has some effect that, however, cannot be detected in this analysis (Figure 2D) suggesting furin engagement at later steps in the infection cascade. This is in line with earlier in vitro studies on the role of cell surface furin (Day et al., 2008; Day and Schiller, 2009; Surviladze et al., 2015). In any case, our results align with both models of ECM detachment: one involving HS cleavage (HS co-transfer) and another one involving capsid modification (by e.g., KLK8).

Earlier version: We propose that after 5 h of CytD treatment, glycan-induced structural activation as well as capsid processing by proteases such as KLK8 of the capsid and HS cleavage essentially are completed ((i) in Figure 9). Subsequently, (ii) the HS decorated virion is transported from the ECM to the cell body on filopodia, which takes about 15 min. In a time window of 30 min, PsV-CD151 association occurs. (iii) In the samples recorded 30 – 180 min after CytD wash off, we observe that PsVs lose their HS-coat, and individual PsV-CD151 assemblies seem to merge into larger structures that are subsequently endocytosed. This model highlights the role of ECM interactions, actin dynamics, and receptor engagement in HPV16 entry.

Revised version: We propose that after 5 h of CytD treatment, glycan-induced structural activation as well as capsid processing by proteases such as KLK8 of the capsid and HS cleavage essentially are completed ((i) in Figure 9). Subsequently, the HS decorated virion is recruited from the ECM to the basal cell membrane, which takes about 15 min. In a time window of 30 min after CytD wash off (ii), PsV-CD151 association occurs. (iii) In the samples recorded 30 – 180 min after CytD wash off, we observe that PsVs lose their HS-coat, and individual PsV-CD151 assemblies seem to merge into larger structures that are subsequently endocytosed. This model highlights the role of ECM interactions, actin dynamics, and receptor engagement in HPV16 entry.

Materials and Methods

Antibodies and PsVs

We used the following primary antibodies in immunostainings of proteins. For the capsid protein L1, a rabbit polyclonal antibody (pAb) K75 (diluted 1:1000) as described previously (Knappe et al., 2007; Rommel et al., 2005), for CD151 a mouse monoclonal antibody (mAb) (1:200) (Biorad, cat# MCA1856GA), and for Itgα6 a rabbit polyclonal antibody (1:200) (Invitrogen, cat# PA5-12334) were used. For Heparan sulfate (HS) a mouse IgM monoclonal antibody (1:200) (amsbio, cat# 370255-S) was used that reacts with an epitope in native heparan sulfate chains and not with hyaluronate, chondroitin or DNA, and poorly with heparin (mAb 10E4 (David et al., 1992)). For HS neo-epitope (Yokoyama et al., 1999) detection, a mouse monoclonal antibody (1:200) (amsbio, cat#370260-S) was used that reacts only with heparitinase-treated heparan sulfate chains, proteoglycans, or tissue sections, and not with heparinase treated HSPGs. The antibody recognizes desaturated uronic acid residues (mAb 3G10 (David et al., 1992)). As secondary antibodies we used a STAR GREEN-coupled goat-anti-rabbit (1:200) (Abberior, cat# STGREEN-1002-500UG), an AlexaFluor 594-coupled donkey-anti-rabbit (1:200) (Invitrogen, cat# A21207), an AlexaFluor 594-coupled donkey-anti-mouse (1:200) (Invitrogen, cat# A21203), an AlexaFluor 594-coupled donkey-anti-mouse IgM (1:200) (Invitrogen, cat# A21044), a STAR RED-coupled goat-anti-mouse (1:200) (Abberior, cat# STRED-1001-500UG) and a STAR RED-coupled goat-anti-rabbit (1:200) (Abberior, cat# STRED-1002-500UG). Additionally, phalloidin iFluor488 (1:1000 of ready to use solution) (abcam, cat# ab176753) or phalloidin iFluor647 (1:1000) (abcam, cat# ab176759) was used to stain F-actin. Staining of the PsV plasmid DNA was performed by click-labeling according to the manufacturer’s instructions with the dye 6-FAM Azide (Baseclick EdU 488 kit, Carl Roth, cat# 1Y67.1).

HPV16 PsVs were produced following established procedures (Buck et al., 2004; Spoden et al., 2012). HEK293TT cells were cultured in 175 cm² flasks and transfected with polyethylenimine (PEI), using equimolar amounts of codon-optimized HPV16 L1/L2 (pShell-16L1wt-16L2wt) and a reporter plasmid (pGL4.20-puro-HPV16 LCR (Wüstenhagen et al., 2018)). For the production of 5-ethynyl-2’-deoxyuridine (EdU)-labeled PsVs, the culture medium was replaced 5 hours post-transfection with fresh medium containing 30 μM EdU. After 48 hours, cells were harvested, centrifuged, and washed twice with phosphate-buffered saline (PBS) supplemented with 9.5 mM MgCl₂ (1x PBS/ MgCl2). After the final centrifugation, the pellet was resuspended in 1x PBS/ MgCl2 containing 0.5% Brij58 (Sigma-Aldrich), and 250 units Benzonase (Merck Millipore), and incubated at 37°C for 24 hours on a rotating platform. Subsequently, lysates were chilled on ice and the NaCl concentration was adjusted to 0.85 M. After clarification by centrifugation, the supernatant was loaded onto an iodixanol (Optiprep) gradient consisting of 39%, 33%, and 27% layers (bottom to top). Gradients were equilibrated for 90 minutes at room temperature before ultracentrifugation at 55,000 rpm for 3.5 hours at 16°C. Fifteen fractions of 300 µl each were collected from the top and analyzed via luciferase reporter assay to identify peak fractions. PsV titers were determined based on packaged genomes (viral genome equivalents, vge) as previously described (Spoden et al., 2012). The concentrations of stock solutions were 7.7 x 106 vge/µl (for microscopy experiments) and 14.1 x 106 vge/µl (for the luciferase assay).

Cell culture

For microscopy, human immortalized keratinocytes (HaCaTs) (Cell Lines Services, Eppelheim, Germany) were cultured in high glucose Dulbecco’s modified Eagle’s medium (DMEM + GlutaMAX, Gibco, cat# 61965-026) supplemented with 10% FBS (PAN Biotech, cat# P30-3031) and 1% Penicillin/Streptomycin (10,000 U/ml Penicillin, 10 mg/ml Streptomycin; PAN Biotech, cat# P06-07100) at 37 °C with 5% CO2. For the luciferase infection assay, cells were grown at 37°C and 5% CO2 in Dulbecco’s modified Eagle’s medium (DMEM + GlutaMAX, Thermo Fisher Scientific), supplemented with 10% FBS (Sigma-Aldrich), 1% minimum essential medium non-essential amino acids (MEM non-essential amino acids (Thermo Fisher Scientific)), and 5 µg/ml ciprofloxacin (Fresenius Kabi). For experiments, the antibiotics in the medium were omitted.

Sample preparation and immunostaining

Earlier version: About 150,000 HaCaT cells were plated onto 25 mm diameter poly-L-lysine (PLL) coated (100 µg/ml PLL for 30 min) glass coverslips in 6-well plates and incubated for 24 h at 37°C and 5% CO2. The next day, cells were incubated at 37°C and 5% CO2 for 5 h with PsVs (50 vge/plated cell) and 10 µg/ml cytochalasin D (CytD; stock solution 10 mg/ml in dimethyl sulfoxide (DMSO)) (Life Technologies, cat# PHZ1063) or 30 µM (−)-blebbistatin (stock solution 13.68 mM in DMSO) (Sigma-Aldrich, cat# B0560-1MG) in DMEM supplemented with 10% FBS. For controls, we added the same amount of DMSO without CytD or blebbistatin. Cells were washed with PBS (137 mM NaCl, 2.7 mM KCl, 1.76 mM KH2PO4, 10 mM Na2HPO4, pH 7.4) and fresh DMEM supplemented with 10% FBS was added. Cells were incubated further at 37°C and 5% CO2 for 0 min (here the medium was added followed by immediate removal), 15 min, 30 min, 60 min and 180 min.

Revised version: About 150,000 HaCaT cells were plated onto 25 mm diameter poly-L-lysine (PLL) coated (100 µg/ml PLL for 30 min) glass coverslips in 6-well plates and incubated for 24 h at 37°C and 5% CO2. The next day, cells were incubated at 37°C and 5% CO2 for 5 h with PsVs (46 vge/plated cell) and 10 µg/ml cytochalasin D (CytD; stock solution 10 mg/ml in dimethyl sulfoxide (DMSO)) (Life Technologies, cat# PHZ1063) or 30 µM (−)-blebbistatin (stock solution 13.68 mM in DMSO) (Sigma-Aldrich, cat# B0560-1MG) in DMEM supplemented with 10% FBS. For controls, we added the same amount of DMSO without CytD or blebbistatin. Cells were washed with PBS (137 mM NaCl, 2.7 mM KCl, 1.76 mM KH2PO4, 10 mM Na2HPO4, pH 7.4) and fresh DMEM supplemented with 10% FBS was added. Cells were incubated further at 37°C and 5% CO2 for 0 min (here the medium was added followed by immediate removal), 15 min, 30 min, 60 min and 180 min.

Earlier version: In another set of experiments, about 150,000 HaCaT cells were plated onto 25 mm diameter PLL coated glass coverslips in 6-well plates and incubated for 24 h at 37°C and 5% CO2. The next day, cells were incubated at 37°C and 5% CO2 for 5 h with PsVs (50 vge/plated cell) and 10 µg/ml CytD either together with 100 µM leupeptin (stock solution 100 mM in ddH2O) (Carl Roth, cat# CN33.1) or 5 µM Furin inhibitor I (stock solution 5 mM in DMSO) (Sigma-Aldrich, cat# 344930-1MG) in DMEM supplemented with 10% FBS. For controls, we used only 10 µg/ml CytD. Cells were washed with PBS and fresh DMEM supplemented with 10% FBS was added. Cells were incubated further at 37°C and 5% CO2 for 0 min, 30 min and 60 min.

Revised version: In another set of experiments, about 150,000 HaCaT cells were plated onto 25 mm diameter PLL coated glass coverslips in 6-well plates and incubated for 24 h at 37°C and 5% CO2. The next day, cells were incubated at 37°C and 5% CO2 for 5 h with PsVs (46 vge/plated cell) and 10 µg/ml CytD either together with 100 µM leupeptin (stock solution 100 mM in ddH2O) (Carl Roth, cat# CN33.1) or 5 µM Furin inhibitor I (stock solution 5 mM in DMSO) (Sigma-Aldrich, cat# 344930-1MG) in DMEM supplemented with 10% FBS. For controls, we used only 10 µg/ml CytD. Cells were washed with PBS and fresh DMEM supplemented with 10% FBS was added. Cells were incubated further at 37°C and 5% CO2 for 0 min, 30 min and 60 min.

Earlier version: To analyze PsV binding to detached HaCaT cells, about 350,000 HaCaT cells per well were plated in 6-well plates and incubated for 24 h at 37°C and 5% CO2. The next day, cells were detached by a 15 min incubation with 10 mM EDTA (in PBS, pH 7.4) at 37°C and 5% CO2. Detached cells were collected and incubated with 50 vg/cell in DMEM supplemented with 10% FBS under constant rotation at 4 °C for 1 h. Cells were washed three times with PBS at 4 °C to remove any unbound PsVs. Then, cells were seeded onto 25 mm diameter PLL coated glass coverslips in 6-well plates and incubated for 1 h at 37°C and 5% CO2.

Revised version: To analyze PsV binding to detached HaCaT cells, about 350,000 HaCaT cells per well were plated in 6-well plates and incubated for 24 h at 37°C and 5% CO2. The next day, cells were detached by a 15 min incubation with 10 mM EDTA (in PBS, pH 7.4) at 37°C and 5% CO2. Detached cells were collected and incubated with 46 vge/cell in DMEM supplemented with 10% FBS under constant rotation at 4 °C for 1 h. Cells were washed three times with PBS at 4 °C to remove any unbound PsVs. Then, cells were seeded onto 25 mm diameter PLL coated glass coverslips in 6-well plates and incubated for 1 h at 37°C and 5% CO2.

Before staining, cells were washed twice with PBS and fixed at room temperature (RT) with 4% PFA in PBS for 30 min, unless membrane sheets were generated. In this case, the coverslips were placed in ice cold sonication buffer (120 mM KGlu, 20 mM KAc, 10 mM EGTA, 20 mM HEPES, pH 7.2) and a 100 ms ultrasound pulse at 100% power was applied. This was repeated until in total about 10 pulses were applied at different locations of the coverslip. Then, membrane sheets were fixed like cells at RT with 4% PFA in PBS for 30 min. PFA was removed, and residual PFA was quenched by 50 mM NH4Cl in PBS for 30 min. Afterwards, samples were blocked with 3% BSA in PBS for 30 min. Staining of PsVs was performed by click-labeling with the dye 6-FAM for 30 min at RT according to the manufacturer’s instructions. Then, samples were washed three times with PBS. In case of no PsV labeling by click-chemistry, directly after blocking, the respective primary antibodies were added: mouse IgM mAb against HS (1:200), mouse mAb against Δ-HS (1:200), rabbit pAb against Itgα6 (1:200), mouse mAb against CD151 (1:200) and rabbit pAb K75 against L1 (1:1000) in 3% BSA in PBS for 2 h. Samples were washed three times with PBS before adding the respective secondary antibodies and fluorescent labelled phalloidins in 3% BSA in PBS for 1 h: for HS we added AlexaFluor 594 coupled to donkey-anti-mouse IgM (1:200), for Δ-HS AlexaFluor 594 coupled to donkey-anti-mouse (1:200), for Itgα6 STAR RED coupled to goat-anti-rabbit (1:200), for CD151 AlexaFluor 594 coupled to donkey-anti-mouse (1:200) or STAR RED coupled to goat-anti-mouse (1:200), for L1 AlexaFluor 594 coupled to donkey-anti-rabbit (1:200) or STAR GREEN coupled to goat-anti-rabbit (1:200) and for F-actin phalloidin iFluor488 (1:1000) or phalloidin iFluor647 (1:1000). Afterwards, samples were washed three times with PBS, followed by mounting of the coverslips onto microscopy slides using ProLong Gold antifade mounting medium (Invitrogen, cat# P36930).

Confocal and STED microscopy

Earlier version: For confocal and STED microscopy, samples were imaged employing a 4-channel STED microscope from Abberior Instruments. The microscope is based on an Olympus IX83 confocal microscope equipped with an UPlanSApo 100x (1.4 NA) objective (Olympus, Tokyo, Japan). Confocal and STED micrographs were recorded simultaneously. For confocal imaging, a 488 nm laser was used for the excitation of 6-FAM, STAR GREEN and phalloidin iFluor488, and emission was detected at 500– 550 nm. For STED imaging, a 561 nm laser (at 45%) was used for the excitation of AlexaFluor 594 (detection at 580–630 nm) and a 640 nm laser (at 45%) for the excitation of STAR RED and iFluor647-labelled phalloidin (detection at 650-720 nm), in combination with a 775 nm laser for depletion (at 45%). The pixel size was set to 25 nm, for overview images a pixel size of 250 nm was used.

Revised version: For confocal and STED microscopy, samples were imaged employing a 4-channel STED microscope from Abberior Instruments (available at the superresolution light microscopy facility of the LIMES institute, Bonn, Germany). The microscope is based on an Olympus IX83 confocal microscope equipped with an UPlanSApo 100x (1.4 NA) objective (Olympus, Tokyo, Japan). Confocal and STED micrographs were recorded simultaneously. For confocal imaging, a 488 nm laser was used for the excitation of 6-FAM, STAR GREEN and phalloidin iFluor488, and emission was detected at 500–550 nm. For STED imaging, a 561 nm laser (at 45%) was used for the excitation of AlexaFluor 594 (detection at 580–630 nm) and a 640 nm laser (at 45%) for the excitation of STAR RED and iFluor647-labelled phalloidin (detection at 650-720 nm), in combination with a 775 nm laser for depletion (at 45%). The pixel size was set to 25 nm, for overview images a pixel size of 250 nm was used.

Analysis of confocal and STED micrographs

Images were analyzed using Fiji ImageJ (Schindelin et al., 2012) in combination with a custom written macro, essentially as described previously (Schmidt et al., 2024). When analyzing double stainings of Alexa594 (red channel) and STAR RED (long red channel), we corrected for crosstalk from the red into the long red channel by subtracting 40% of the intensity of the red channel image from the long red channel image.

First, images were smoothed with a Gaussian blur to improve maxima detection. For confocal PsV-DNA micrographs, we employed a Gaussian blur of σ = 1, and for confocal PsV-L1 and STED micrographs a Gaussian blur of σ = 0.5. Local maxima were detected using the ‘Find Maxima’ function (noise tolerance 60 for L1 if stained together with phalloidin iFluor488, and 80 if stained together with CD151, HS, or phalloidin iFluor647, 3 for click-labeled PsVs, 8 for CD151, 6 for HS and 15 for Itgα6), yielding maxima positions in pixel positions. For analysis, regions of interest (ROIs) were defined (for details see figure legends).

In order to define the cell border region (Supplementary Figure 3B), the ImageJ ‘Make Binary’ function was used on CD151 STED micrographs for generating a binary mask. When possible, the ‘Wand’ tool was used to further outline the area covered by the cell, otherwise the same was done manually using the original micrograph as a reference. From the outline in the binary mask, a ROI was created, that was then filled in white with the ’Flood Fill’ Tool. The ROI was shrunk to a 20-pixel broad strip, and cleared outside of the strip, what leaves a 20-pixel broad, closed ribbon marking the intracellular side of the cell border that, however, exhibits arbitrary edges produced in the process above. These edges were removed by manual adjustments with the ’Pencil’ tool and the clear function. Afterwards, the ROI defines the intracellular side of the cell border. The ROI was symmetrically broadened by 40 pixels (using the ‘Enlarge’ function). From the now at least 60 pixels broad ROI, approximately two thirds covered the intracellular and one third the extracellular side. In this ROI, using the ImageJ ‘Find Maxima’ function, PsV-L1 maxima were detected (at a noise tolerance of 80) with a custom written macro in ImageJ both in the cell border region ROI and in a ROI covering the entire micrograph (excluding a 2-pixel edge). From these values, the percentage of PsVs within the cell border region was calculated.

For measuring intensity over time, we measured in the ROI the mean gray value. For background correction, we subtracted the mean gray value measured in a ROI next to the cell.

For measuring maxima intensity, a ROI with a diameter of 125 nm (5 pixels) was placed onto the determined maxima positions (see above). Using these ROIs, the average mean gray value of each maximum was measured. The average background mean gray value was measured in a ROI placed next to the cell and subtracted from the average mean gray value of the maxima. For each membrane sheet, the average maxima intensity was calculated, followed by averaging of the values of the membrane sheets.

To measure the shortest distance, e.g., between PsV and CD151 maxima, as a quality control for the PsV maxima (CD151 maxima do not undergo the quality control step) we placed onto each maximum position a horizontal and a vertical linescan (31 pixels long × 3 pixels width). Afterwards, we fitted a Gaussian distribution to the intensity distribution of each maximum. Only if at least one of the fits exhibits a fit quality of R2 > 0.8 and if the Gaussian maximum located central to the intensity distribution the maximum was included in the further analysis. Moreover, using a 125 nm ROI at the PsV and CD151 maxima positions, the center of mass of fluorescence was determined, yielding the maxima positions in sub-pixels. Based on these positions we further calculated the shortest distance of a PsV maximum to the nearest CD151 maximum.

Using the shortest distances, we either calculated the fraction of closely associated PsVs (PsVs with a distance ≤ 80 nm to e.g., CD151) over time, or the average shortest distance over time. The fraction of closely associated PsVs was corrected for random background association that increases with the maxima density. The relationship between background association and maxima density was expressed by a linear equation obtained from the same images re-analyzed after flipping them horizontally and vertically (see e.g., Supplementary Figure 7).

Earlier version: The Pearson correlation coefficient (PCC) was calculated with a custom written macro in ImageJ. As a control, we performed the analysis on horizontally and vertically flipped ROI-defined images as well (see e.g., Supplementary Figure 6A). The cell body PCC between PsVs and HS is measured in smaller ROIs covering exclusively the cell body region (this is confirmed with the Itgα6 or F-actin image as reference).

Revised version: The Pearson correlation coefficient (PCC) was calculated with a custom written macro in ImageJ. As a control, we performed the analysis on horizontally and vertically flipped ROI-defined images as well (see e.g., Supplementary Figure 5A). The cell body PCC between PsVs and HS is measured in smaller ROIs covering exclusively the cell body region (this is confirmed with the Itgα6 or F-actin image as reference).

Per condition and biological replicate, 14 - 15 images were analyzed and values were averaged. For images from one set of experiments, in the figure the same channels are shown at the same settings of brightness and contrast, if not stated otherwise.

Epi-fluorescence microscopy and image analysis

For epifluorescence microscopy, we used microscopic equipment and settings as previously described (Homsi et al., 2014) except of the illumination system, which was replaced by a SPECTRA X Light engine; Lumencor, Beaverton, OR, USA). In brief, PFA-fixed cells were imaged in a chamber filled with 1 ml of PBS to which 50 µl of a saturated solution of TMA-DPH [1-(4-tri-methyl-ammonium-phenyl)-6-phenyl-1,3,5-hexatriene-p-toluenesulfonate (T204; Thermofisher)] in PBS was added, for visualizing the membranes in the blue channel. PsVs, HS and Itgα6 were imaged in the green, red and far-red channels, respectively.

For analysis, rectangular ROIs of ≈ 45 x 103 pixel were placed onto the cell body such that approximately one half of the ROI covers the ECM area and the other one the cell body. The cell periphery was defined using the TMA-DPH image from which a binary mask was created. With reference to the binary mask, a band of up to 30 pixels width was created beginning at the cell body and reaching out towards the cell periphery. Using this ROI, in the PsV-DNA image the size of the ROI and the mean gray values were measured from which the integrated intensities were calculated after background correction.

Luciferase infection assay

HaCaT cells were grown in 24-well plates and allowed to adhere overnight. The following day, cells were treated with HPV16 PsVs at a concentration of approximately 100 vge per cell, in the presence or absence of 10 µg/ml cytochalasin D (CytD) in DMSO or an equivalent amount of DMSO (Control). In one condition, PsVs/CytD were applied for 5 hours, after which the medium was replaced with fresh medium lacking the compounds, and incubation was continued for additional 19 hours (in total 24 h). In another condition, cells were exposed to PsVs/CytD continuously for the full 24 h period. Then, cells were washed with PBS and lysed using 1 x Cell Culture Lysis Reagent (Promega, cat# E153A). Following high-speed centrifugation, luciferase activity in the cleared lysates was quantified using an LB 942 Tristar 3 Multimode Microplate Reader (Berthold Technologies). Cytotoxic effects, accompanied by a loss of membrane integrity (indicated by released lactate dehydrogenase (LDH)), were determined by measuring the LDH levels in the cell lysates. LDH activity was assessed using the CytoTox-ONE Homogeneous Membrane Integrity Assay (Promega, cat# G7891). LDH fluorescence was quantified according to the manufacturer’s instructions using the LB 942 Tristar 3 Multimode Microplate Reader.

Statistics

Microscopy data sets were based on three biological replicates. One replicate includes per time point the average of 14 - 15 images analyzed (with exception of Figure 4, which includes 14 – 35 analyzed cells per replicate and time point). For the infection assay, data sets were based on three biological replicates. One biological replicate is the average of three technical replicates. Data was tested for significance with a two-tailed, unpaired student’s t-test with significance * = p < 0.05, ** = p < 0.01 and *** = p < 0.001.